```python
# Initialize Otter
import otter
grader = otter.Notebook("lab02.ipynb")
```

# [Functions] Lab 02 – Iteration and Algorithms

## Data 6, Fall 2025

Welcome to Lab 10! In this lab, we'll be covering **iteration** and **algorithms**.

```python
# Just run this cell
from datascience import *
import numpy as np
import warnings
warnings.simplefilter('ignore')
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

# [Tutorial] Part 1: Iteration with `for` Loops
Iteration is the repeated running of code. Typically, the variables in iteration change sequentially, and the iteration continues to repeat until a boolean condition indicates that the iteration be stopped.

One way to do this is using a `for` loop. A `for` loop begins with the word `for`, followed by a name we want to give each item in the sequence, followed by the word `in`, and ending with an expression that evaluates to a sequence. The indented body of the `for` loop is executed once for each item in that sequence.

For example, say we're given the following array, where each value corresponds to the number of hours of sleep one night of the week: 

```python
hours_of_sleep = make_array(6, 7, 5, 8, 7, 6, 8)
```

We could use a `for` loop to _iterate_ through each value in the array. One simple thing we could do is print each value in the array. 

```python
for num_hours in hours_of_sleep: 
    print(num_hours)
```

Notice that we used `num_hours` as our name for each item in the sequence, but we could have given it another name, and it wouldn't change the functionality of the `for` loop: 

```python
for n in hours_of_sleep: 
    print(n)
```

While the above `for` loop is fairly simple, we could do more than just printing each value. For example: 

```python
night_of_sleep = 1
for num_hours in hours_of_sleep: 
    print("On night " + str(night_of_sleep)+ ", I slept for " + str(num_hours) + " hours.")
    night_of_sleep += 1
```

In the code above, we "initialize" `night_of_sleep` to 1, and use it to "counter" to keep track of what night we are on.

<br></br>
<hr style="border: 1px solid #fdb515;" />

# Question 1 – FizzBuzz

[FizzBuzz](https://en.wikipedia.org/wiki/Fizz_buzz) is a game and common programming challenge. It is described below.

> Players generally sit in a circle. The player designated to go first says the number "1", and the players then count upwards in turn. However, any number divisible by three is replaced by the word "fizz" and any number divisible by five by the word "buzz". Numbers divisible by 15 become "fizz buzz". A player who hesitates or makes a mistake is eliminated from the game.

In this problem, you will write a function that computes the "FizzBuzz sequence." We'll practice using `for` loops to do this.

---
## Question 1.1

First, complete the implementation of the function `single_number`, which takes in an integer (`n`) greater than or equal to 1 and returns the corresponding number or word in the FizzBuzz sequence, **as a string**. For instance, `single_number(2)` should evaluate to `'2'`, and `single_number(30)` should evaluate to `'fizz buzz'`.

_Hint_: Use the remainder operator `%`.<br>
_Hint_: You might find conditionals useful.


```python
def single_number(n):
    ...
```

```python
grader.check("q1_1")
```

---
## Question 1.2

Complete the implementation of the function `fizz_buzz_sequence`, which takes in an integer (`n`) greater than or equal to 1 and returns an array of length `n` containing the first `n` elements of the FizzBuzz sequence.

For instance, `fizz_buzz_sequence(8)` should evaluate to the array `array(['1', '2', 'fizz', '4', 'buzz', 'fizz', '7', '8'])`.

_Hint_: Use `np.append`. For more information on `np.append`, you can reference the `NumPy` documentation [here](https://numpy.org/doc/stable/reference/generated/numpy.append.html). 


```python
def fizz_buzz_sequence(n):
    sequence = make_array()
    for i in np.arange(1, n+1, 1):
        ...
    return sequence
```

```python
grader.check("q1_2")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

# [Tutorial] Part 2: Strings

We have already seen strings before in this course, but we want to introduce some new ideas with strings that will be very useful when dealing with data from real-world data sets.

Strings can be viewed as special lists of characters. As we learned in HW04, strings can be printed, sliced, indexed, and more.

One important property about strings is that they are **immutable**, meaning their values cannot be changed once created. For example, say we have the following string: 

```python
our_string = 'flag'
print(our_string)
```

Notice that the following code will result in an error:

```python
our_string[3] = 'x'
```

Instead, we would have to create a new string, using other approaches such as slicing and string concatenation.

```python
our_new_string = our_string[:3] + 'x'
print(our_new_string)
```

Another thing we could do with strings is check if it's _in_ (or _not in_) an array. 

```python
data_classes = make_array('Data 6', 'Data 8', 'Data 100')
"Data 6" in data_classes
```

```python
"Data 8" in data_classes
```

```python
"Data 5" not in data_classes
```

---
## Question 2.1

Complete the implementation of the function `only_uppercase`, which takes in a string `s` and returns a new string containing only the uppercase characters in `s`. Example behavior is shown below.

```py
>>> only_uppercase('University of California, BerkelEY')
'UCBEY'

>>> only_uppercase('now THiS!')
'THS'
```

*Hint*: Use `.isupper()` on each character individually. Confused about what `.isupper()` will do? Try checking its documentation by Googling the function name.


```python
def only_uppercase(s):
    output = ''
    for char in s:
        ...
    return output
```

```python
grader.check("q2_1")
```

---
## Question 2.2

Let's use this idea of array and string indexing and `if/else` to find all words that begin with a vowel. Complete the function `starts_with_vowels` so that it takes an array of strings, finds all of the words that begin with a vowel, and then returns an array of said words.

For the example array, `some_words`, your function should return `array(['alligator', 'aardvark', 'owl', 'unicorn'])`.

*Hint*: You'll need to use the `np.append` function; here is the [documentation]("https://numpy.org/doc/stable/reference/generated/numpy.append.html").

```python
some_words = make_array("alligator", "monkey", "zebra", "lion", "aardvark", "owl", "bear", "unicorn")

def starts_with_vowels(words):
    words_to_keep = make_array()
    for ... in ...:
        if ... in "aeiou":
            ...
starts_with_vowels(some_words)
```

```python
grader.check("q2_2")
```

---
## Question 2.3

Let's write a function that uses another `for` loop. We want to be able to find the acronym of a name. For example, we want to be able to call **The University of California, Berkeley** by an abbreviated `UCB`.

Because we don't care as much about small words even if they are capitalized, also make sure to **exclude words with three letters or fewer from your acronym**. That is why the `T` in `The` is not included in the example acronym above.

Write a function that takes in a string and returns its acronym.

_Hint_: the string method `.isupper()` will return `True` if the given string (or character) is entirely upper case, and `False` otherwise.

```python
# Here is a cell that explains the str.split() function used in the solution below:
sentence = "Hello, I am a student in Data 6!"
sentence.split()
```

```python
def acronym(title):
    # This str.split() function mentioned above turns a string into a list of its words.
    words = title.split()
    result = ""
    ...
    return result

acronym("The University of Califiornia, Berkeley")
```

```python
grader.check("q2_3")
```

---
## Question 3: Hailstone (revisited)
You might remember the Hailstone sequence from Lab 9! As a reminder, the sequence from Hofstadter is as follows:
1. Pick a positive integer n that is greater than 1 as the start value.
2. If n is even, divide it by 2.
3. If n is odd, multiply it by 3 and add 1.
4. If you continue this process, n will eventually reach 1. 

In Lab 9, we wrote a function that completes **one step** of the hailstone sequence. This time, we'll be writing a function that runs the complete sequence, which prints each value of n, and **returns the number of iterations run**. <br>
_Note_: Skeleton code is provided, you are not required to follow it. 


```python
def hailstone(n): 

    count = ...
    ...
    while ...: 
        ...
        if ...:
            ...
        else: 
            ...
        print(n)
```

```python
grader.check("q3")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

# Part 3: Comparisons with NumPy
Learning `np.any()` and `np.all()` is essential since it allows much faster operations in comparison to iterating over the entire array.

Here is the [np.all() documentation](https://numpy.org/doc/stable/reference/generated/numpy.all.html), and [np.any() documentation](https://numpy.org/doc/stable/reference/generated/numpy.any.html). 

The function `np.all()` takes in an array and tests whether all array elements evaluate to truthy values. Run the following cells and verify that the output matches what you expect. 

np.all(make_array(True, True, False))

```python
np.all(make_array(1, 1, 1, 1))
```

```python
np.all([1, 2.5, -1])
```

Meanwhile, `np.any()` takes in an array and tests whether any array element is evaluated as True. 

```python
np.any(make_array(1, 0))
```

```python
 np.any(make_array(False, False, False, False))
```

```python
# The following line of code errors due to a syntax error. Are you able to spot it?
np.any(make_array'fizz, 'buzz'))
```

# Question 4
How are `np.any()` and `np.all()` equivalent to compound OR and compound AND?


_Type your answer here, replacing this text._

----
## Done! 🤠

## Pets of Data 6

Portobello says congrats on completing Lab 10!

<img src="portobello_caramello.png" width="50%" alt="Cute pets"/>

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q1_1 (points: 1)
```
>>> callable(single_number)
True
```
```
>>> single_number(30) == 'fizz buzz'
True
```
```
>>> single_number(9) == 'fizz'
True
```
```
>>> single_number(10) == 'buzz'
True
```
```
>>> # Make sure you're returning a string, not an int!
>>> single_number(1) == '1'
True
```


Tests for q1_2 (points: 1)
```
>>> all(fizz_buzz_sequence(8) == make_array('1', '2', 'fizz', '4', 'buzz', 'fizz', '7', '8'))
True
```
```
>>> all(fizz_buzz_sequence(3) == make_array('1', '2', 'fizz'))
True
```
```
>>> fizz_buzz_sequence(15).item(14) == 'fizz buzz'
True
```
```
>>> sum(fizz_buzz_sequence(1000) == "fizz") == 267
True
```


Tests for q2_1 (points: 1)
```
>>> callable(only_uppercase)
True
```
```
>>> only_uppercase('University of California, BerkelEY') == 'UCBEY'
True
```
```
>>> only_uppercase('now THiS!') == 'THS'
True
```
```
>>> only_uppercase('Data 6, Summer 2021') == 'DS'
True
```
```
>>> only_uppercase('what if they are all lowercase?') == ''
True
```


Tests for q2_2 (points: 1)
```
>>> all(starts_with_vowels(some_words) == make_array('alligator', 'aardvark', 'owl', 'unicorn'))
True
```
```
>>> all(starts_with_vowels(make_array("bear")) == make_array())
True
```
```
>>> "zebra" not in starts_with_vowels(some_words)
True
```


Tests for q2_3 (points: 1)
```
>>> callable(acronym)
True
```
```
>>> acronym("The University of California, Berkeley") == "UCB"
True
```
```
>>> acronym("") == ""
True
```
```
>>> acronym("Data Science at Berkeley") == "DSB"
True
```
```
>>> acronym("The United States of America") == "USA"
True
```


Tests for q3 (points: 1)
```
>>> callable(hailstone)
True
```
```
>>> hailstone(2) == 1
2
1
True
```
```
>>> # verify that hailstone returns the count, not the value of n
>>> hailstone(3) != 10 and hailstone(10) != 5
3
10
5
16
8
4
2
1
10
5
16
8
4
2
1
True
```
```
>>> hailstone(9) == 19 and hailstone(10) == 6
9
28
14
7
22
11
34
17
52
26
13
40
20
10
5
16
8
4
2
1
10
5
16
8
4
2
1
True
```
