```python
# Initialize Otter
import otter
grader = otter.Notebook("Homework_1.ipynb")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

# Homework 1: Built-in Functions & Variables
Version: August 20th, 2025

This homework assignment focuses on 3 aspects of coding:

1. Primitive Data Types
2. Variables
3. Built-in Functions



<hr style="border: 1px solid #fdb515;" />

## Part A: Variables

To get familiar with variables, we are going to create a profile that contains information about you through storing information through variables.

<div class="alert alert-info">
<b>Question 1:</b> Let's make our profile by assigning values to variables! For refreshers, you can look back on our lecture content.
</div>

```python
# Fill out the ... part with your information.
# If you do not fill comfortable sharing your profile, you can make up a profile.

name = ...
age = ...
city = ...
state = ...

# These variables are checking whether you are a student/have a job right now, what data type would work best?
student = ...
job = ...
```

<div class="alert alert-info">
<b>Question 2:</b> Use the print() built-in function to output the values assigned to the variables.
</div>

<details>
    <summary>Sample Output: </summary>
    <img src="images/Q2_edited.png" width="1000"/>
</details>


```python
#Type your answer here
```

<hr style="border: 1px solid #fdb515;" />

## Part B: Data Type

Now let's try to understand data types of these variables. For each question, assign your answer from the multiple choice in a string format.

        Example) What is a cat?
        1.   A feline species
        2.   A human
        3.   An alien

        answer = "1"


<div class="alert alert-info">
<b>Question 3:</b> What is the data type of `name`?

  1.   Integers
  2.   Floats
  3.   Strings
  4.   Booleans
  5.   None
</div>

```python
answer_3 = ...
```

```python
grader.check("q3")
```

<div class="alert alert-info">
<b>Question 4:</b> What is the data type of `age`?

  1.   Integers
  2.   Floats
  3.   Strings
  4.   Booleans
  5.   None
</div>

```python
answer_4 = ...
```

```python
grader.check("q4")
```

<div class="alert alert-info">
<b>Question 5:</b> List all the variables that are strings:

<ul>
    <li>name</li>
    <li>age</li>
    <li>city</li>
    <li>state</li>
    <li>student</li>
    <li>job</li>
</ul>

*Type your answer here*

<div class="alert alert-info">
<b>Question 6:</b> Where do you see yourself in 5 years? Let's create a profile of your future self.

To do this, fill in where you see yourself in 5 years into the new variables listed below.
</div>

```python
five_year_later_local = ...
five_year_later_job = ...
```

<div class="alert alert-info">
Next, use the print() built-in function to output the values assigned to the variables.
</div>
<details>
    <summary>Sample Output: </summary>
    <img src="images/Q6_edited.png" width="1000"/>
</details>

```python
#Type your answer here
```

<hr style="border: 1px solid #fdb515;" />

## Part C: Arithmetic Operators

In python, we can use arithmetic operators to increase the value for age, meaning it isn't necessary to create a new variable!

Make sure to use the previous variables from Question 1 in your updated profile print statement.

<div class="alert alert-info">
<b>Question 7:</b> Repeat question 6 but include your age after 5 years in the updated portion of your printed statement.
</div>

<details>
    <summary>Sample Output: </summary>
    <img src="images/Q7_edited.png" width="1000"/>
</details>


```python
#Type your answer here
```

<hr style="border: 1px solid #fdb515;" />

## Part D: Lists and Sorting

An integral part of python programming is utilizing **lists** to store information, whether it be numerical or categorical.

<div class="alert alert-info">
<b>Question 8:</b> You are having a garage sale, and you need to keep track of the <b>clothes</b> you have that need to be sold. 
    
Create a list named <b>clothes</b>, which has the following items: <i>shirts, pants, shorts, jackets, and hats</i>

```python
clothes = ...
```

```python
grader.check("q8")
```

<div class="alert alert-info">
<b>Question 9:</b> Now, create anohter list called <b>inventory</b> which keeps track of just the number of each clothing item you have: 7 shirts, 4 pants, 2 shorts, 5 jackets, and 3 hats.

```python
inventory = ...
```

```python
grader.check("q9")
```

<div class="alert alert-info">
<b>Question 10:</b> How many total items of clothing are you selling?

<i>Hint: Use the sum() function!</i>
</div>

```python
total = ...
print(total)
```

```python
grader.check("q10")
```

<div class="alert alert-info">
<b>Question 11:</b> Output the sorted inventory list using a built-in function.
    
(Sort the inventory list by name alphabetically)
</div>

```python
sorted_list = ...
print(sorted_list)
```

```python
grader.check("q11")
```

<div class="alert alert-info">
<b>Question 12:</b> Find the Mean for the Frozen Desserts Data Set

Output the Mean for the Grams of Sugar Column Only using Arithmetic operators (+,-,*,/)

Properly format the print() function to label the answer
</div>

<img src="images/frozen_desserts.jpg" style="width:50%"/>


```python

```

```python
dessert_mean = ...
```

```python
grader.check("q12")
```

<div class="alert alert-info">
<b>Question 12:</b> Find the Mean for the Sugar Bomb Data Set

Output the Mean for the Grams of Sugar Column Only using Arithmetic operators (+,-,*,/)

Properly format the print() function to label the answer
</div>

<img src="images/sugar_bomb.png" style="width:50%" />

```python
sbomb_mean = ...
```

```python
grader.check("q13")
```

```python

```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q10 (points: 2)
```
>>> isinstance(total, int)
True
```
```
>>> total == 21
True
```


Tests for q11 (points: 2)
```
>>> isinstance(sorted_list, list)
True
```
```
>>> len(sorted_list) == 5
True
```
```
>>> sorted_list == ['hats', 'jackets', 'pants', 'shirts', 'shorts']
True
```


Tests for q12 (points: 2)
```
>>> type(dessert_mean) == float
True
```
```
>>> dessert_mean - 20.2 < 0.0001
True
```


Tests for q13 (points: 2)
```
>>> type(sbomb_mean) == float
True
```
```
>>> sbomb_mean - 22.77 < 0.0001
True
```


Tests for q3 (points: 2)
```
>>> type(answer_3) == type('word')
True
```
```
>>> answer_3 == '3'
True
```


Tests for q4 (points: 2)
```
>>> type(answer_4) == type('word')
True
```
```
>>> answer_4 == '1'
True
```


Tests for q8 (points: 2)
```
>>> isinstance(clothes, list)
True
```
```
>>> len(clothes) == 5
True
```


Tests for q9 (points: 2)
```
>>> isinstance(inventory, list)
True
```
```
>>> len(inventory) == 5
True
```
