```python
# Initialize Otter
import otter
grader = otter.Notebook("Homework_3.ipynb")
```

```python
# Don't change this cell; just run it. 

import numpy as np
import pandas as pd
import matplotlib.pyplot as plt
%matplotlib inline
#import d8error
```

# Homework 3: Pandas

The main topics covered in Homework 2 are:
<ul>
    <li> .csv / Data Frames</li>
    <li> Pandas</li>
</ul>

---

## Part 1: Pandas

For this next part of the homework, we will be using Numpy and python's pandas library to analyze real-world data sets.

We will be looking at <b>Individual Hospital Statistics for California.</b> These are non-federal, short-term, acute care hospitals in California. The dataset can be accessed directly using the following link or by the website:
* [California Hospital Website](https://www.ahd.com/states/hospital_CA.html)

Here are the data sets for both the SF and LA Hospitals:
* [Hospitals_LA](https://docs.google.com/spreadsheets/d/1ZWb0wUIu-IV1QM8G6pHNJUHrz4rAtnYxb5aX6M-iVe0/edit?usp=sharing)
* [Hospitals_SF](https://docs.google.com/spreadsheets/d/1VSN0961WQQGbrAugjI5-iN5FnsmYIjWYYz-tR-wWbqo/edit?usp=sharing)

Note: You will already have these data sets accessible within the notebook with the names of the data sets above.

Use the Pandas LIbrary, Hospitals_SF.csv, and Hospitals_LA.csv to complete the questions below. 


<b>San Francisco Hospitals (13):</b>
<ul>
    <li>Chinese Hospital</li>
    <li>CPMC Davies Campus</li>
    <li>CPMC Mission Bernal Campus & Orthopedic Institute</li>
    <li>CPMC Van Ness Campus</li>
    <li>Kaiser Permanente San Francisco Medical Center</li>
    <li>Laguna Honda Hospital and Rehabilitation Center</li>
    <li>San Francisco VA Medical Center</li>
    <li>UCSF Betty Irene Moore Women's Hospital</li>
    <li>UCSF Health Hyde Hospital</li>
    <li>UCSF Health Stanyan Hospital</li>
    <li>UCSF Helen Diller Medical Center at Parnassus Heights</li>
    <li>UCSF Mission Bay Campus</li>
    <li>Zuckerberg San Francisco General Hospital and Trauma Center</li>
</ul>

<b>Los Angeles Hospitals (18):</b>
<ul>
    <li>Adventist Health White Memorial</li>
    <li>California Hospital Medical Center</li>
    <li>Cedars-Sinai Medical Center</li>
    <li>CHA Hollywood Presbyterian Medical Center</li>
    <li>Docs Surgical Hospital</li>
    <li>East Los Angeles Doctors Hospital</li>
    <li>LA Downtown Medical Center Downtown Campus</li>
    <li>Los Angeles Community Hospital</li>
    <li>Los Angeles General Medical Center</li>
    <li>Los Angeles Medical Center</li>
    <li>Martin Luther King, Jr. Community Hospital</li>
    <li>PIH Health Good Samaritan Hospital</li>
    <li>Ronald Reagan UCLA Medical Center</li>
    <li>The Keck Hospital of USC</li>
    <li>UCLA Mattel Children's Hospital</li>
    <li>USC Norris Comprehensive Cancer Center</li>
    <li>West Los Angeles Medical Center</li>
    <li>West Los Angeles VA Medical Center</li>
</ul>

---

### Sum methods - Pandas

This section will focus on using Pandas `sum()` methods directly on our data sets.

<div class="alert alert-info">
    <b>1. Use the Pandas .sum() method to calculate total Staffed_Beds for San Francisco Hospitals.</b>
</div>

```python
...
staffed_beds_SF = ...
print(staffed_beds_SF)
```

```python
grader.check("q1")
```

<div class="alert alert-info">
    <b>2. Use the Pandas .sum() method to calculate total Staffed_Beds for Los Angeles Hospitals.</b>
</div>

```python
...
staffed_beds_LA = ...
print(staffed_beds_LA)
```

```python
grader.check("q2")
```

<div class="alert alert-info">
    <b>3. Use the Pandas .sum() method to calculate total Total_discharges for San Francisco Hospitals.</b>
</div>

```python
...
total_discharges_SF = ...
print(total_discharges_SF)
```

```python
grader.check("q3")
```

<div class="alert alert-info">
    <b>4. Use the Pandas .sum() method to calculate total Total_discharges for Los Angeles Hospitals.</b>
</div>

```python
...
total_discharges_LA = ...
print(total_discharges_LA)
```

```python
grader.check("q4")
```

<div class="alert alert-info">
    <b>5. Use the Pandas .sum() method to calculate total Gross_Patient_Revenue for San Francisco Hospitals.</b>
</div>

```python
...
gpr_SF = ...
print(gpr_SF)
```

```python
grader.check("q5")
```

<div class="alert alert-info">
    <b>6. Use the Pandas .sum() method to calculate total Gross_Patient_Revenue for Los Angeles Hospitals.</b>
</div>

```python
...
gpr_LA = ...
print(gpr_LA)
```

```python
grader.check("q6")
```

---
### Mean & Median Methods - Pandas

This section will focus on using Pandas `mean()` method directly on our data sets.

<div class="alert alert-info">
    <b>7. Find the mean for total Gross_Patient_Revenue for San Francisco Hospitals.</b>
</div>

```python
...
mean_gpr_SF = ...
print(mean_gpr_SF)
```

```python
grader.check("q7")
```

<div class="alert alert-info">
    <b>8. Find the mean for total Gross_Patient_Revenue for Los Angeles Hospitals.</b>
</div>

```python
...
mean_gpr_LA = ...
print(mean_gpr_LA)
```

```python
grader.check("q8")
```

<div class="alert alert-info">
    <b>9. Find the median for total Gross_Patient_Revenue for San Francisco Hospitals.</b>
</div>

```python
...
median_gpr_SF = ...
print(median_gpr_SF)
```

```python
grader.check("q9")
```

<div class="alert alert-info">
    <b>10. Find the median for total Gross_Patient_Revenue for Los Angeles Hospitals.</b>
</div>

```python
...
median_gpr_LA = ...
print(median_gpr_LA)
```

```python
grader.check("q10")
```

<div class="alert alert-info">
    <b>11. Compare and contrast the mean number of staffed beds for San Francisco and Los Angeles Hospitals. Are the results the same results you received for homework 2?
</b>
</div>

*Type answer here*

---
### Visualization Methods - Histograms

In this section, we will be using Pandas and Matplotlib to graph histograms of certain columns of our data sets.

<div class="alert alert-info">
    <b>12. Create a Histogram for the Gross_Patient_Revenue variable/column in the data set for San Francisco Hospitals.

[Hints*** panda and matplotlib example]

</b>
</div>

```python
# Type answer here

```

<div class="alert alert-info">
    <b>13. Looking at the histograms Gross_Patient_Revenue. Estimate the typical revenue for hospitals in San Francisco. Is the mean value you calculated above similar to your estimate based upon the histogram? 
</b>
</div>

*Type answer here*

<div class="alert alert-info">
    <b>14. Create a Histogram for the Gross_Patient_Revenue variable/column in the data set for Los Angeles Hospitals.

[Hints*** panda and matplotlib example]

</b>
</div>

```python
# Type answer here
```

<div class="alert alert-info">
    <b>15. Looking at the histograms Gross_Patient_Revenue. Estimate the typical revenue for hospitals in Los Angeles Hospital.  Is the mean value you calculated above similar to your estimate based upon the histogram? 
</div>

*Type answer here*

```python

```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q1 (points: 2)
```
>>> type(staffed_beds_SF) == np.int64
True
```
```
>>> staffed_beds_SF - pd.read_csv('Hospitals_SF.csv', sep=',')['Staffed_Beds'].sum() < 1e-05
True
```


Tests for q10 (points: 2)
```
>>> type(median_gpr_LA) == float
True
```
```
>>> median_gpr_LA - pd.read_csv('Hospitals_LA.csv', sep=',')['Gross_Patient_Revenue'].median() < 1e-05
True
```


Tests for q2 (points: 2)
```
>>> type(staffed_beds_LA) == np.int64
True
```
```
>>> staffed_beds_LA - pd.read_csv('Hospitals_LA.csv', sep=',')['Staffed_Beds'].sum() < 1e-05
True
```


Tests for q3 (points: 2)
```
>>> type(total_discharges_SF) == np.float64
True
```
```
>>> total_discharges_SF - pd.read_csv('Hospitals_SF.csv', sep=',')['Total_Discharges'].sum() < 1e-05
True
```


Tests for q4 (points: 2)
```
>>> type(total_discharges_LA) == np.float64
True
```
```
>>> total_discharges_LA - pd.read_csv('Hospitals_LA.csv', sep=',')['Total_Discharges'].sum() < 1e-05
True
```


Tests for q5 (points: 2)
```
>>> type(gpr_SF) == np.int64
True
```
```
>>> gpr_SF - pd.read_csv('Hospitals_SF.csv', sep=',')['Gross_Patient_Revenue'].sum() < 1e-05
True
```


Tests for q6 (points: 2)
```
>>> type(gpr_LA) == np.int64
True
```
```
>>> gpr_LA - pd.read_csv('Hospitals_LA.csv', sep=',')['Gross_Patient_Revenue'].sum() < 1e-05
True
```


Tests for q7 (points: 2)
```
>>> type(mean_gpr_SF) == np.float64
True
```
```
>>> mean_gpr_SF - pd.read_csv('Hospitals_SF.csv', sep=',')['Gross_Patient_Revenue'].mean() < 1e-05
True
```


Tests for q8 (points: 2)
```
>>> type(mean_gpr_LA) == np.float64
True
```
```
>>> mean_gpr_LA - pd.read_csv('Hospitals_LA.csv', sep=',')['Gross_Patient_Revenue'].mean() < 1e-05
True
```


Tests for q9 (points: 2)
```
>>> type(median_gpr_SF) == float
True
```
```
>>> median_gpr_SF - pd.read_csv('Hospitals_SF.csv', sep=',')['Gross_Patient_Revenue'].median() < 1e-05
True
```
