```python
# Initialize Otter
import otter
grader = otter.Notebook("project1.ipynb")
```

# Project 1 – Education, Admissions, and Simpson’s Paradox

## Data 6, Fall 2025

### Due Date: Wednesday, October 22 at 5pm

You must submit the assignment to Gradescope. Submission instructions can be found at the bottom of this notebook. See the [syllabus](http://data6.org/fa25/syllabus/#late-policy-and-extensions) for our late submission policy.

**Note:** Unlike in previous homework assignments, most questions in this assignment will depend on all previous work. As such, it's in your best interest to work through the questions sequentially.

### Project Partners

For projects, you can work with up to one partner. Please see our syllabus for our [partner collaboration policy](https://data6.org/fa25/syllabus/#projects).

If you work with a partner, please **include their name** below. **Additionally, submit exactly one copy of the assignment** by listing your partner in your Gradescope submission.

**Partners**:
* <partner 1 name, partner 1 email>
* <partner 2 name, partner 2 email>

## Assignment Setup

In this project, you will explore how visualizations, distributions, and disaggregations can reveal underlying trends in the data. We will focus on UC Berkeley admissions, past and present.

```python
# Run this cell.
from datascience import *
import numpy as np
import matplotlib.pyplot as plt
plt.style.use("ggplot")
%matplotlib inline

import warnings
warnings.simplefilter('ignore')
```

**Note:** It's a good idea to have the [Python Reference Sheet](https://data6.org/su24/reference/) open while working on the assignment in the event you have any questions.

You can also easily see the documentation for a function by either:
- typing the name of the function on a new line, followed by a `?`, and running the cell
- typing the name of the function anywhere in a code cell and hitting `Shift + Tab` on your keyboard

Try it out below! Just add a `?` to the end of the line.

```python
Table.where
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

# [Present] Part 1: UC Berkeley Undergraduate Admissions, 2024

<br></br>
<hr style="border: 1px solid #fdb515;" />

## Question 1: Unboxing the dataset

In this part of the assignment, we explore UC Berkeley's undergraduate admissions numbers for the freshman Fall 2024 class. The data we'll work with in this question comes from [this public webpage](https://www.universityofcalifornia.edu/infocenter/admissions-source-school).

Run the cell below to load in our data as a Table, `schools_full`.

```python
schools_full = Table.read_table('firstyear_admissions_2024.csv')
schools_full
```

Each row corresponds to a high school. For each high school, we have the following information:
- `'School'`: The name of the high school.
- `'City'`: The city the high school is located in.
- `'Region'`: The county the high school is located in if the high school is in California, **or** the state in which the high school is if the high school is elsewhere in the US.
- `'Applied'`: The number of first-year students who applied to UC Berkeley from that high school for admission that year.
- `'Admitted'`: The number of first-year students who were admitted to UC Berkeley from that high school for admission that year.
- `'Enrolled'`: The number of first-year students who actually chose to attend UC Berkeley from that high school that year.

Applied, Admitted, and Enrolled numbers are all with respect to Fall 2024.

---

## Question 1a

How many schools are in our dataset? Use `schools_full`.

```python
num_schools = ...
num_schools
```

```python
grader.check("q1a")
```

Many schools seem to have the same name. For instance, the name 'ABRAHAM LINCOLN HIGH SCHOOL' corresponds to five different high schools: in Los Angeles, San Francisco, Santa Clara in California (CA); and in Council Bluffs and Des Moines, Iowa (IA):

```python
schools_full.where("School", "ABRAHAM LINCOLN HIGH SCHOOL")
```

---

## Question 1b(i)

Create a table `count_by_school_name` that counts the number of schools for each school name. Your table should have the following format:

| School | count | 
| --- | --- |
| ... | ... |
| ABRAHAM LINCOLN HIGH SCHOOL | 5 | 
| ... | ... |

```python
count_by_school_name = ...
count_by_school_name
```

```python
grader.check("q1b_i")
```

---

## Question 1b(ii)

Using `count_by_school_name`, count the number of schools that share a name with at least one other school in our dataset, and assign this to `num_same_name`. In other words, count the number of schools that **do not** have a unique name in our dataset.

_Hint_: How can you filter rows from a table with `where`? How does `count_by_school_name` tell you if a school's name is unique?

```python
num_same_name = ...
num_same_name
```

```python
grader.check("q1b_ii")
```

Next, let's examine the city names. 

---

## Question 1c(i)

What is the most common city in the dataset? Create a table `count_by_city` that counts the number of schools in each city **and** region, in order of **most schools first**. Your table should look like this (here are the first few rows):

| City |	Region |	count |
| --- | --- | --- |
| nan | nan | 2477 |
| Los Angeles | Los Angeles | 173 | 
| New York | NY | 92 |
| ... | ... | ... |

_Hint_: See the [Data 6 Python Reference](https://data6.org/notes/reference) and course notes for how to group rows by multiple attributes.

```python
count_by_city = ...
count_by_city
```

```python
grader.check("q1c_i")
```

It seems like the most common city is called `'nan'`. `'nan'` is a special value in NumPy that means "Not a Number". In many datasets, it is used to represent a "missing value."

Here are the schools that have a missing City and Region:

```python
# just run this cell
schools_full.where("City", "nan").where("Region", "nan")
```

<!-- BEGIN QUESTION -->

---

## Question 1c(ii)

Given the table above, do some research about where these `"nan"` schools are actually located. They all have one common factor related to their geographic location. In 1-2 sentences below, explain what sorts of schools have missing cities and regions in our dataset.

_Hint_: UC Berkeley is a public university in California, United States.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<br></br>
<hr style="border: 1px solid #fdb515;" />

# Question 2 – Admission Rates

<br></br>

---
## Question 2a – How many students were admitted?

We use the term "first-year" or "freshman" to define undergraduate students who are expected to spend eight undergraduate semesters at UC Berkeley—this is in contrast to "junior" or "junior transfer" students, who begin their fifth semester at UC Berkeley.

How many first-year students were **admitted** to UC Berkeley in Fall 2024? Use `schools_full` and assign your answer to `num_admitted`.

```python
num_admitted = ...
num_admitted
```

```python
grader.check("q2a")
```

---
## Question 2b – What was the overall admission rate?

The **admission rate** to a school is defined as the proportion of students who were admitted of those who applied:


$$ \text{admission rate} = \frac{\text{\# admitted}}{\text{\# applied}} $$

Below, use `schools_full` to assign the variable `admission_rate` to the overall admission rate for UC Berkeley first-year students. You may use `num_admitted`, which you defined earlier.


```python
admission_rate = ...
admission_rate
```

```python
grader.check("q2b")
```

---
## Question 2c – Double-checking the data

The admission rate you computed in the previous part should seem peculiarly high.

### Question 2c(i)

Find the true admission rate from this August 2024 [UC Berkeley News](https://news.berkeley.edu/2024/08/02/almost-80-of-uc-berkeleys-first-year-admits-are-from-california/) article. Assign `true_admission_rate_choice` to an integer 1, 2, 3, or 4 corresponding to the choices below. This question has a hidden test.

1. 79%
2. 11%
3. 75%
4. 12%

```python
true_admission_rate_choice = ...
```

```python
grader.check("q2c_i")
```

<!-- BEGIN QUESTION -->

### Question 2c(ii)

_Why is the admission rate computed in our dataset less than the true admission rate as reported by the news article?_ Write a short answer in the cell below. Your answer should include (1) the details listed in the "fine print" on the UC website we used to obtain this data (https://www.universityofcalifornia.edu/about-us/information-center/admissions-source-school), and (2) why these schools and applicants may have been excluded.

_Note_: We used the "FR Gnd by Yr" table to generate our dataset; in this table, a student "category" is defined as All, Female, Male, Other, and Unknown.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

---

## Question 2d: The `schools` Table

Given the above, for the rest of this assignment we will consider only the schools in `schools_full` that **had _admits_ to UC Berkeley for the Fall 2024 admissions cycle**. This restricts the questions we can ask about our data, but we will see that there are still some interesting results.

Assign this new table to `schools`. The `schools` table should have the same columns as `schools_full` but include only the rows with non-zero admits.

```python
schools = ...
schools.show(10)
```

```python
grader.check("q2d")
```

<br></br>
<hr style="border: 1px solid #fdb515;" />

# Question 3 – Location

Next, let's explore student enrollments by location.

---

## Question 3a - Regions

Use the table `schools` to create a new table `regions` which sums up applicant, admission, and enrollment numbers by region.

Your table should look like the following, with a few sample rows:

| Region | Applied | Admitted | Enrolled | 
| --- | --- | --- | --- |
| AK | 66 | 40 | 0 | 
| ... | ... | ... | ... |
| Alameda | 7890 | 6163 | 2931 |
| ... | ... | ... | ... |

_Hints_: Use `group`. Note that the Applied, Admitted, and Enrolled numbers are the totals of schools in that region. You may need to rename/drop columns. You should feel free to write your code in multiple lines or one single line; our code uses two lines.

```python
regions = ...
regions = Table().with_columns(
    "Region", regions.column("Region"),
    "Applied", regions.column("Applied sum"),
    "Admitted", regions.column("Admitted sum"),
    "Enrolled", regions.column("Enrolled sum")
)
regions
```

```python
grader.check("q3a")
```

---

## Question 3b - Which region sent the most students to UC Berkeley?

Assign `top_5_regions` to an **array** of the **region names** of the five regions with the most first-year enrollments to UC Berkeley in Fall 2024. The first element of `top_5_regions` has the most enrollees, the second element, has the second most enrollees, and so on.

You can verify your answer with the [UC Berkeley News](https://news.berkeley.edu/2024/08/02/almost-80-of-uc-berkeleys-first-year-admits-are-from-california/) article we shared earlier. This question has hidden tests.

_Hint_: You will need to sort `regions` in descending or ascending order; the order you choose will impact how you `take` rows.

```python
...
top_5_regions = ...

# Don't change anything below this comment, it's just for display purposes
print('Top 5 regions:')
for region in top_5_regions:
    print(region)
```

```python
grader.check("q3b")
```

<br></br>
<hr style="border: 1px solid #fdb515;" />

# Question 4 – Admission Rate by School

Given the caveats we examined in the previous part, let us first look at schools that had _high_ admission rates.

---

## Question 4a

Make a new table `schools_admit` that adds an `'Admission Rate'` column to `schools` with the admission rate for each school, each as a decimal between 0 (no students were admitted) and 1 (all students were admitted).

You may find it helpful to compute `admission_rate_array` as an array of admission rates per school, then include that as the `'Admission Rate'` column in `schools_admit`.

The first few rows of `schools_admit` might look like this:

| School	| City	| Region	| Applied	| Admitted	| Enrolled	| Admission Rate |
| --- | --- | --- | --- | --- | --- | --- |
| 21ST CENTURY EXPERIMENTAL SCH	| nan	| nan	| 125	| 104	| 33	| 0.832 |
| A B MILLER HIGH SCHOOL	| Fontana	| San Bernardino	| 81 | 	67	| 27	| 0.82716 |
| ... | ... | ... | ... | ... | ... | ... |



```python
admission_rate_array = ...
schools_admit = ...
schools_admit
```

```python
grader.check("q4a")
```

<!-- BEGIN QUESTION -->

---

## Question 4b - Distribution of Admission Rates

### Question 4b(i)

Use `schools_admit` to make a scatter plot of school admission rates by number of admits.

To avoid **overplotting**, you can pass an optional argument `alpha` to adjust the _opacity_ of each point in the scatter plot. `alpha` takes in a fraction, 0 (completely transparent) to 1.0 (completely opaque).

For example, to set half-opacity for a scatter plot of **y by x**, where `y` and `x` are columns in a table `tbl`:

    tbl.scatter(x, y, alpha=0.5)

```python
...
```

<!-- END QUESTION -->

---

### Question 4b(ii)

Based on your visualization above, is it true that the schools with the most students admitted also have the highest admission rates?

Based on your answer, assign `is_highest_admit` to the appropriate Boolean value of `True` or `False`. This question has hidden tests.

```python
is_highest_admit = ...
```

```python
grader.check("q4b_ii")
```

<!-- BEGIN QUESTION -->

---
## Question 4c

### Question 4c(i)

Produce a histogram that visualizes the distributions of admission rates in `schools_admit`, using the given bins in `my_bins`. Let the area of each bar be the percent of entries in each bin.

Hint: See the [Python reference](https://data6.org/notes/reference) if you're stuck on how to specify bins.

```python
my_bins = np.arange(0, 1.01, .1)

...
```

<!-- END QUESTION -->

---

### Question 4c(ii)

How many schools have an admission rate of [0.6, 0.8)? Assign your answer to `num_schools_60_to_80` below.

_Hint_: Use Boolean predicates to filter the appropriate rows in `schools_admit`, then get the count.

```python
num_schools_60_to_80 = ...
num_schools_60_to_80
```

```python
grader.check("q4c_ii")
```

---
## Question 4d – Bay Area schools

The Bay Area consists of the nine counties `'San Francisco'`, `'San Mateo'`, `'Santa Clara'`, `'Alameda'`, `'Contra Costa'`, `'Solano'`, `'Napa'`, `'Sonoma'`, and `'Marin'`.

<img src='https://upload.wikimedia.org/wikipedia/commons/b/bc/Bayarea_map.png' width=400>


Assign `bay_schools` to a table with the same columns as `schools_admit`, but only with rows corresponding to schools in the Bay Area. The `bay_area_counties` array has been provided for your convenience.


```python
bay_counties = make_array('San Francisco', 'San Mateo', 'Santa Clara',
                          'Alameda', 'Contra Costa', 'Solano',
                          'Napa', 'Sonoma', 'Marin')

bay_schools = ...
bay_schools
```

```python
grader.check("q4d")
```

<br></br>
<hr style="border: 1px solid #fdb515;" />

# Question 5 - Yield Rate by School

The **yield rate** to a school is defined as the proportion of students who enrolled of those who were _admitted_:

$$ \text{yield rate} = \frac{\text{\# enrolled}}{\text{\# admitted}} $$

While the admission rate is often used by prospective students to determine the chances they might get into a school, by contrast the **yield rate** is often used by schools to determine whether the students they are admitting are _actually going to attend_.

---

## Question 5a

Make a new table `schools_yield` that adds a `'Yield Rate'` column to `schools_admit` with the **yield rate** for each school, each as a decimal between 0 (no admitted students enrolled) and 1 (all admitted students enrolled).

You may find it helpful to compute `yield_rate` as an array of yield rates per school, then include that as the `'Yield Rate'` column in `schools_yield`.

The first few rows of `schools_admit` might look like this:

| School	| City	| Region	| Applied	| Admitted	| Enrolled	| Admission Rate	| Yield Rate |
| --- | --- | --- |--- | --- | --- | --- | --- |
| 21ST CENTURY EXPERIMENTAL SCH| nan	| nan	| 125	| 104	| 33	| 0.832	| 0.317308 |
| A B MILLER HIGH SCHOOL	| Fontana	| San Bernardino	| 81	| 67	| 27	| 0.82716	| 0.402985 |
| ... | ... | ... | ... | ... | ... | ... | ... |

```python
yield_rate = ...
schools_yield = ...
schools_yield
```

```python
grader.check("q5a")
```

<!-- BEGIN QUESTION -->

---

## Question 5b - Yield vs. Admission

### Question 5b(i)

Use `schools_yield` to make a scatter plot of school yield rates by admission rates.

To avoid **overplotting**, you can pass an optional argument `alpha` to adjust the _opacity_ of each point in the scatter plot. `alpha` takes in a fraction, 0 (completely transparent) to 1.0 (completely opaque).

For example, to set half-opacity for a scatter plot of **y by x**, where `y` and `x` are columns in a table `tbl`:

    tbl.scatter(x, y, alpha=0.5)

```python
...
```

<!-- END QUESTION -->

---

### Question 5b(ii)

Based on your visualization above, are higher admission rates correlated with higher yield rates? In other words, do schools with higher admission rates tend to have more students who enroll?

Based on your answer, assign `is_admit_yield_correlated` to the appropriate Boolean value of `True` or `False`. This question has hidden tests.

```python
is_admit_yield_correlated = ...
```

```python
grader.check("q5b_ii")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

# [Past] Part 2: UC Berkeley Undergraduate Admissions, 1973

Now, we are going to look at a very interesting dataset, the 1973's UC Berkeley Admission data for data analysis! We will explore and reproduce tables and visualizations in the original Bickel, Hammel, and O'Connell paper from 1975.

References:
* Bickel, Hammel, and O’Connell. “Sex Bias in Graduate Admissions:” Data from Berkeley“ _Science_ vol. 187 (1975) pp. 398-404. [Google Drive link, Fall 2025](https://drive.google.com/file/d/1LXCJ1NzIr01WUJMyqez2MBY0YgDKZEly/view?usp=drive_link)
* "Simpson's Paradox." Wikipedia. https://en.wikipedia.org/wiki/Simpson%27s_paradox
* "Berkeley's 1973 Graduate Admissions Dataset." Data Science Discovery, University of Illinois Urbana-Champaign. https://discovery.cs.illinois.edu/dataset/berkeley/
* Freedman, Pisani, and Purves. “Sex Bias in Graduate Admissions.” _Statistics_ (1978) Chapter 2.  pp. 17-20.

<!-- BEGIN QUESTION -->


<hr style="border: 1px solid #fdb515;" />

# Question 6: Understanding the Case Study

Let's begin exploring this paradox by going through the Bickel et al. paper ([Google Drive link, Fall 2025](https://drive.google.com/file/d/1LXCJ1NzIr01WUJMyqez2MBY0YgDKZEly/view?usp=drive_link)). We'll get an idea of what steps they took in exploring the data they had from the 1973 graduation admissions, and we'll analyze some of their findings and figures to see how they connect to the paradox.

From the "Data and Assumptions" section on the first page, we see that the authors had data on all approximately 15,000 applications. Their main purpose of exploring this data was to see if the sex of the applicant in any way influenced the decision to admit or deny any particular student. If there was any statistically significant result from this exploration, there would be a bias on sex present, and from there, they might hope to explore if a discrimination also existed.

---

## Question 6a

In your own words, explain the authors' first approach in determining whether a sex bias was present, as described on page 1. Using this method, what results did they find? Does there appear to be a sex bias when using this first method? From your own judgment (and also information from the two paragraphs about the assumptions made, if you'd like), do you think this method is the most effective in exploring if a sex bias is present?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

In the "Some Underlying Dependencies" section on page 2, the authors bring up how their second assumption (the assumption that "the sex ratios of applicants to the various fields of graduate study are not importantly associated with any other factors in admission") being false results in a case of Simpson's paradox. They describe it as:

> We have stumbled onto a paradox, sometimes referred to as Simpson's in this context (1) or 'spurious correlation' in others (2). It is rooted in the falsity of assumption 2 above. We have assumed that if there is a bias in the proportion of women applicants admitted it will be because of a link between sex of applicant and decision to admit. We have given much less attention to a prior linkage, that between sex of applicant and department to which admission is sought.

---
## Question 6b

From the statement above, which of the following options describes an important, initially overlooked factor that the authors believe is associated with the sex of the applicant? Assign `overlooked_factors` to an array of integers that correspond to all valid options below, e.g., `make_array(3, 4)`.

1. Department the applicant applies to
2. Whether the student was an athlete or not
3. Decision to admit the applicant or not
4. Race/ethnicity of the student

```python
overlooked_factors = ...
```

```python
grader.check("q6b")
```

## Simpson's Paradox
According to [Wikipedia](https://en.wikipedia.org/wiki/Simpson%27s_paradox), Simpson's Paradox "is a phenomenon in probability and statistics in which a trend appears in several groups of data but disappears or reverses when the groups are combined."

At first glance of the UC Berkeley Admission dataset, it seems male have higher chance being admitted. However, when looked closer by department, a different story emergered. We found that females were equal or more likely to be admitted when looked at each department. In fact, the original Bickel, Hammel, and O'Connell study found that female applicants disproportionately applied to highly competitive departments with low admission rates (e.g., English), while men more often applied to less competitive departments with higher admission rates. *More female applicants applying to competitive programs skewed the combined data, creating the appearance of bias against female applicants, when in many cases, women were either equally likely or more likely to be admitted compared to men in the same department.*

## Exploring the Dataset

Now that we've explored some of the early findings from the original paper, let us look into exploring Simpson's Paradox by reproducing some of the figures in the paper—and making some of our own visualizations.

A portion of the original 1973 dataset has been recreated in the `grad_admission_1973.csv` file, which we load into the `cal_data` table. This includes the admission numbers for the largest six departments and groups other departments together in an "Other" category.

For the remainder of this assignment, you may find it helpful to keep a copy of the original paper handy: [Google Drive link, Fall 2025](https://drive.google.com/file/d/1LXCJ1NzIr01WUJMyqez2MBY0YgDKZEly/view?usp=drive_link)

```python
# just run this cell
cal_data = Table.read_table("grad_admission_1973.csv")
cal_data
```


<hr style="border: 1px solid #fdb515;" />

# Question 7: Admission Rates by Sex

## Question 7a: `admit_numbers`

Assign the table `admit_numbers` to one that contains the counts of female and male applicants admitted and denied in the `cal_data` table. Your table should have the following shape, though your rows may be in a different order:

| Sex | Admit | Deny |
| --- | --- | --- |
| F | ... | ... |
| M | ... | ... |

_Hint_: Use `pivot`.

```python
admit_numbers = ...
admit_numbers
```

```python
grader.check("q7a")
```

---

## Question 7b: `admit_rates`

Next, make a new table `admit_rates` that extends `admit_numbers` by two new columns:
* `"Applied"` which computes the number of applicants by sex
* `"Admission Rate"` which computes the admission rate by sex

We have created the arrays `admitted` and `denied` for the two corresponding columns in `admit_numbers`.

Your new table `admit_rates` should have the following shape, though your rows may be in a different order:


| Sex | Admit | Deny | Applied | Admission Rate |
| --- | --- | --- | --- | --- |
| F | ... | ... | ... | ... |
| M | ... | ... | ... | ... |

_Hint_: You may find it helpful to create the array `applied` which you will then assign as the new `"Applied"` column.

```python
# provided for you
admitted = admit_numbers.column("Admit")
denied = admit_numbers.column("Deny")

# complete the below code
applied = ...
admit_rates = ...
admit_rates
```

```python
grader.check("q7b")
```

<!-- BEGIN QUESTION -->

## Question 7c

Using the adjusted `admit_rates`, create a bar graph comparing the acceptance rate between female and male applicants. If you are stuck, consider taking a look at Lab 3 again!

```python
...
```

<!-- END QUESTION -->

## Question 7d

Which of the following most accurately describes the takeaway from the visualization created by `admit_rate` above? You should answer the question by assigning `bar_observation` to an integer 1, 2, or 3 corresponding to the correct choice below.

1. Female applicants have a higher likelihood of being admitted.
2. Male applicants have a higher likelihood of being admitted.
3. Female and male applicants have the same likelihood of being admited.

```python
bar_observation = ...
```

```python
grader.check("q7d")
```

<!-- BEGIN QUESTION -->

## Question 7e
A media company wants to make a report on UC Berkeley admissions. What would happen if they ONLY report with the visualization you created in this part? How would this report potentially impact who applies to UC Berkeley graduate school?

_Type your answer here, replacing this text._

<!-- END QUESTION -->


<hr style="border: 1px solid #fdb515;" />

# Question 8: Reproducing Table 1

We would like to reproduce Table 1 in the original paper. 

This question uses `admit_rates` above, so please verify your numbers above before continuing.

```python
# just run this cell
admit_rates
```

---

## Question 8a: Overall admission rate

### Question 8a(i)
Compute `admission_rate_overall`, the overall admission rate across both sexes listed in `admit_rates`. Do not round your answer.

_Hint_: You may find it useful to first compute `num_admitted_overall` and `num_denied_overall`, the overall counts of applicants `admitted` and `denied`, respectively.

```python
num_admitted_overall = ...
num_denied_overall = ...

admission_rate_overall = ...
admission_rate_overall
```

```python
grader.check("q8a_i")
```

---

### Question 8a(ii)

Is the female admission rate _lower_ than, _higher_, than or _comparable to_ the overall admission rate? Assign `female_admit_compare` to one of the strings `"lower"`, `"higher"`, or `"comparable"`. This question has hidden tests.

```python
female_admit_compare = ...
```

```python
grader.check("q8a_ii")
```

## Expected Number of Admissions and Denials

The _expected_ number of admits for a particular group is computed based on the two assumptions in the paper—assuming that both sexes are admitted based on the overall admission rate. In other words, for a particular set of applicants:

<!--
$$\text{admit expected} = \text{applied} \times \text{overall admission rate} $$

$$\text{deny expected} = \text{applied} \times (1 - \text{overall admission rate}) $$
-->

The overall "denial" rate is determined so that the total number of expected admitted and denied students equals the number of applied students.

---
## Question 8b: Compute Table 1
    
With this, we are ready to make a version of Table 1 that has the following shape:

| Sex | Admit | Deny | Admit Expected | Deny Expected |
| --- | --- | --- | --- | --- |
| F | ... | ... | ... | ... | 
| M | ... | ... | ... | ... |

Create `table1` from `admit_rates`, `admission_rate_overall`, and any other tables or arrays you defined in earlier parts. This question may be more involved than previous parts.

_Hints_:

1. You can use the `applied` array defined in earlier parts to both create new columns and check your work.
2. You will need to `drop` columns.
3. To further check your work, look at the numbers listed in Table 1. You do not need to round your answers like in the original paper; your rows may also be ordered differently (e.g., female applicants first).

```python
...
table1 = ...

table1
```

```python
grader.check("q8b")
```

<hr style="border: 1px solid #fdb515;" />

# Question 9: Admission by Department

We see by simply utilizing `"Sex"` and `"Decision"` to calculate overall admission rates for graduate admissions, the male admission rate seems to be higher than the female admission rate. However, beyond `"Sex"` and `"Decision"`, there was an additional column in this table: `"Department"`. Let's investigate how the department will take into play with admission rate.


## Question 9a: `applicants_by_dept`

Use `cal_data` to create the table `applicants_by_dept` to one that contains the counts of female and male applicants **broken down by department**.

Your table should have the following shape. Note the **column labels**.

| Department | F Applied | M Applied |
| --- | --- | --- |
| A | ... | ... |
| B | ... | ... |
| ... | ... | ... |


```python
applicants_by_dept = ...
applicants_by_dept = ...
applicants_by_dept
```

```python
grader.check("q9a")
```

## Question 9b: `admits_by_dept`

Use `cal_data` to create the table `admits_by_dept` to one that contains the counts of female and male applicants **_admitted_, broken down by department**.

Your table should have the following shape. Note the **column labels**.

| Department | F Admit | M Admit |
| --- | --- | --- |
| A | ... | ... |
| B | ... | ... |
| ... | ... | ... |

**Note**: Contrast this with the previous part, which was just looking for the total admits—admitted or denied—to each department. How might you adjust your approach above to include _only_ the rows of admitted applicants and _filter out_ denied applicants? (Consider using method chaining.)

```python
admits_by_dept = ...
admits_by_dept = ...
admits_by_dept
```

```python
grader.check("q9b")
```

## Question 9c

Create a table `numbers_by_dept` that combines the two tables (`applicants_by_dept` and `admits_by_dept`). Your new table should have the following shape:

| Department | F Applied | M Applied | F Admit | M Admit |
| --- | --- | --- | --- | --- |
| A | ... | ... | ... | ... |
| B | ... | ... | ... | ... |
| ... | ... | ... | ... | ... |

```python
numbers_by_dept = ...
numbers_by_dept
```

```python
grader.check("q9c")
```

## Question 9d

Finally, create table `admit_rates_by_dept` that computes the admission rates of each sex, broken down by department. Your new table should have the following shape:

| Department | F Admission Rate | M Admission Rate |
| --- | --- | --- | 
| A | ... | ... | 
| B | ... | ... | 
| ... | ... | ... | 

_Hints_:
* Use the `numbers_by_dept` table you just created.
* You might find it useful to create an empty table with `Table()` first, then add new columns. Many approaches are possible, however!

```python
admit_rates_by_dept = Table().with_columns(
    "Department", numbers_by_dept.column("Department"),
    "F Admission Rate", numbers_by_dept.column("F Admitted")/numbers_by_dept.column("F Applied"),
    "M Admission Rate", numbers_by_dept.column("M Admitted")/numbers_by_dept.column("M Applied")
)
admit_rates_by_dept
```

```python
grader.check("q9d")
```

<!-- BEGIN QUESTION -->

<hr style="border: 1px solid #fdb515;" />

# Question 10: Visualizing by Department

In class, we explored why this particular case study exhibited Simpson's Paradox on https://setosa.io/simpsons/. An animated visualization showed the confounding factor—the department itself, where departments with different applicant pools had admission rates that were more or less competitive.

Bar charts can also help reveal this confounding factor, as we see below.

---
    
## Question 10a

Use one of the tables you created in the previous question to create an overlaid bar chart of the **number of applicants** of males and females across the different departments. Each department should have two bars (corresponding to the numbers of male and female applicants).

```python
...
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

---

## Question 10b

Use one of the tables you created in the previous question to create an overlaid bar chart of the **acceptance rates** of males and females across the different departments. Each department should have two bars (corresponding to the admission rates of male and female applicants).

```python
...
```

<!-- END QUESTION -->

---

## Question 10c

Which of the following multiple choice correctly describes *the visualization created above* with `admission_dept` and `num_applicants`? You should answer the question by assigning `overlaid_bar_observation` to `make_array(...)` where `...` is the choice of your answer (eg. `make_array(3, 4)`). Your choice can be can multiple.

1. Among the 6 different departments (not including `"Other"`, female applicants have a slightly higher chance of being admitted than male applicants in several departments.
2. Among the 6 different departments (not including `"Other"`), male applicants have a slightly higher chance of being admitted than female applicants in several departments.
3. Overall, female applicants have a higher chance of getting admitted to UC Berkeley.
4. Overall, male applicants have a higher chance of getting admitted to UC Berkeley. 
5. There seems to be **higher female applicants** to a department when there are **lower female admission rate** compared to male admission rate.  
6. With the visualization above, we can claim there is structural inequity against female women applicants. 

```python
overlaid_bar_observation = ...
```

```python
grader.check("q10c")
```

<!-- BEGIN QUESTION -->

## Question 10d

Consider the takeaways that you had after looking at the different visualizations you created in the above parts. Based on what we've discussed regarding the various sociological paradigms, how do you think these conclusions best align with the paradigms?

*Hint:* The conclusions might not line up perfectly with any one of the paradigms! Just feel free to reflect on which ones they might line up best or worst with. 

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

<hr style="border: 1px solid #fdb515;" />

# [open-ended] Question 11: Reproducing Figure 1

As our last task in this project, let's work towards reproducing Figure 1 in the original paper:  ([Google Drive link, Fall 2025](https://drive.google.com/file/d/1LXCJ1NzIr01WUJMyqez2MBY0YgDKZEly/view?usp=drive_link))

Figure 1 graphs the **percents of applicants** admitted by the **percent of female applicants**, by department. The size of each datapoint encodes the relative number of applicants admitted.

* The **percent admitted** is the proportion of total admits to total applicants.
* The **percent female applicants** is the proportion of female applicants to total applicants.


In the cell below, write code that reproduces the Figure 1 scatter plot to the best of your ability. We've kept this question as **open-ended as possible** for you to explore various approaches. 

**Notes**: With the tools we have in this class _and_ the dataset we are provided, any visualization you create will have a few limitations:

* Our dataset only includes the six largest departments (and groups together other departments into a seventh "Other"). Plot only the **six** largest departments (filter out the "Other") department.
* We do not have the ability to graph squares using our `scatter` Table method. However, you can provide the optional argument `sizes` to adjust the size of each marker. This should be a string corresponding to the name of the numeric column to use for sizes.
* You do not need to plot the 'line of best fit' shown in the original Figure 1.

<details>
        <summary><b>HINTS: See our staff solution outline and resulting plot</b></summary>

Our staff took the following approach to this problem (though note that `col1`, `col2`, and `col3` are placeholders for what should be more meaningful names in your solution):
* Create a new table, e.g., `fig1_table`, that has the relevant columns. This new table referenced and used columns from tables created in previous parts.
    * First, add a new column `col3` for the total number of applicants (summing up  male and female applicants) by department.
    * Then, make two new columns for the admission rate and female application rate by department. Make any necessary additional columns to compute these.
    * Filter out the row corresponding to the "seventh" department, Other.
    * Finally, make columns `Percent women applicants` and `Percent admitted` that convert these two rates to percentages (out of 100) and relabel these columns to prepare for plotting.
* Call `fig1_table.scatter('Percent women applicants', 'Percent admitted', sizes='Total Admitted')`
        
<img src='fig1_staff.png' width=500>

</details>

```python
# be as creative as you want for this problem!
...
```

# Done!

Congrats! You've finished your first Data 6 project!

This project is out of **64 points**. The point breakdown for this assignment is given in the table below:

| **Category** | Points |
| --- | --- |
| Autograder | 44 |
| Written | 20 |
| **Total** | 64 |

---

## Pets of Data 6
Sophia is proud of you :-) congrats on completing Project 1!

<img src="sophia.jpeg" width="50%" alt="Two cats laying on bed"/>

## Submission

Below, you will see two cells. Running the first cell will automatically generate a PDF of all questions that need to be manually graded, and running the second cell will automatically generate a zip with your autograded answers. You are responsible for submitting both the coding portion (the zip) and the written portion (the PDF) to their respective Gradescope portals. **Please save before exporting!**

> **Important: You must correctly assign the pages of your PDF after you submit to the correct gradescope assignment. If your pages are not correctly assigned and/or not in the correct PDF format by the deadline, we reserve the right to award no points for your written work.**

If there are issues with automatically generating the PDF in the first cell, you can try downloading the notebook as a PDF by clicking on `File -> Save and Export Notebook As... -> PDF`. If that doesn't work either, you can manually take screenshots of your answers to the manually graded questions and submit those. Either way, **you are responsible for ensuring your submission follows our requirements, we will NOT be granting regrade requests for submissions that don't follow instructions.**

```python
from otter.export import export_notebook
from os import path
from IPython.display import display, HTML
name = 'project1'
export_notebook(f"{name}.ipynb", filtering=True, pagebreaks=True)
if(path.exists(f'{name}.pdf')):
    display(HTML(f"Download your PDF <a href='{name}.pdf' download>here</a>."))
else:
    print("\n Pdf generation failed, please try the other methods described above")
```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q10c (points: [0])
```
>>> isinstance(overlaid_bar_observation, np.ndarray)
True
```


Tests for q1a (points: 1)
```
>>> answer = schools_full.num_rows
>>> num_schools == answer
True
```


Tests for q1b_i (points: 1)
```
>>> isinstance(count_by_school_name, Table)
True
```
```
>>> [s.lower() for s in count_by_school_name.labels] == [s.lower() for s in ("School", "count")]
True
```
```
>>> answer = 10548
>>> count_by_school_name.num_rows == answer
True
```


Tests for q1b_ii (points: 1)
```
>>> isinstance(num_same_name, int)
True
```
```
>>> answer = 550
>>> num_same_name == answer
True
```


Tests for q1c_i (points: 2)
```
>>> isinstance(count_by_city, Table)
True
```
```
>>> count_by_city.num_rows == 4079
True
```
```
>>> [s.lower() for s in count_by_city.labels] == [s.lower() for s in ("City", "Region", "count")]
True
```
```
>>> count_by_city.column("City").item(0) == "nan"
True
```
```
>>> count_by_city.column("Region").item(0) == "nan"
True
```
```
>>> count_by_city.column("count").item(0) == 2477
True
```


Tests for q2a (points: 1)
```
>>> type(num_admitted) in set([int, np.int32, np.int64, np.float64])
True
```
```
>>> 126000 < num_admitted < 127000
True
```
```
>>> answer = 126132
>>> num_admitted == num_admitted
True
```


Tests for q2b (points: 1)
```
>>> type(admission_rate) in set([float, np.float32, np.float64])
True
```
```
>>> 0.66 < admission_rate < 0.67
True
```
```
>>> np.isclose(admission_rate, 0.66974995619321709)
True
```


Tests for q2c_i (points: 1)
```
>>> isinstance(true_admission_rate_choice, int)
True
```
```
>>> true_admission_rate_choice in [1, 2, 3, 4]
True
```


Tests for q2d (points: 2)
```
>>> sorted(schools.labels) == sorted(schools_full.labels)
True
```
```
>>> schools.num_rows
4346
```
```
>>> schools.column("Applied").sum()
185995.0
```
```
>>> schools.column("Admitted").sum()
126132.0
```


Tests for q3a (points: 2)
```
>>> regions.num_columns
4
```
```
>>> regions.num_rows
111
```
```
>>> sorted(regions.labels) == sorted(('Region', 'Applied', 'Admitted', 'Enrolled'))
True
```
```
>>> regions.where("Region", "Alameda").column("Applied").take(0)
7980.0
```
```
>>> regions.where("Region", "Alameda").column("Enrolled").take(0)
2931.0
```
```
>>> regions.where("Region", "AK").column("Enrolled").take(0)
0.0
```


Tests for q3b (points: 1)
```
>>> isinstance(top_5_regions, np.ndarray)
True
```
```
>>> len(top_5_regions)
5
```


Tests for q4a (points: 2)
```
>>> isinstance(schools_admit, Table)
True
```
```
>>> sorted(schools_admit.labels) == sorted(list(schools.labels) + ['Admission Rate'])
True
```
```
>>> schools_admit.num_rows == schools.num_rows
True
```
```
>>> sorted(schools.labels) == sorted(('School', 'City', 'Region', 'Applied', 'Admitted', 'Enrolled')) # schools doesn't change
True
```
```
>>> subset = schools_admit.where("School", are.contained_in(["21ST CENTURY EXPERIMENTAL SCH", "ABINGTON HEIGHTS HIGH SCHOOL"]))
>>> all(np.isclose(subset.column("Admission Rate"), [0.832, 0.8]))
True
```


Tests for q4b_ii (points: 1)
```
>>> isinstance(is_highest_admit, bool)
True
```


Tests for q4c_ii (points: 1)
```
>>> isinstance(num_schools_60_to_80, int)
True
```
```
>>> num_schools_60_to_80
2078
```


Tests for q4d (points: 2)
```
>>> # Make sure that you're assigning `bay_schools` to a table of Bay Area schools!
>>> isinstance(bay_schools, Table)
True
```
```
>>> sorted(bay_schools.labels) == sorted(('School',
...  'City',
...  'Region',
...  'Applied',
...  'Admitted',
...  'Enrolled',
...  'Admission Rate'))
True
```
```
>>> bay_schools.num_rows
339
```
```
>>> 'ALAMEDA SCIENCE & TECHLGY INST' in bay_schools.column("School")
True
```


Tests for q5a (points: 2)
```
>>> isinstance(schools_yield, Table)
True
```
```
>>> sorted(schools_yield.labels) == sorted(list(schools.labels) + ['Admission Rate', 'Yield Rate'])
True
```
```
>>> schools_yield.num_rows == schools_admit.num_rows == schools.num_rows
True
```
```
>>> sorted(schools.labels) == sorted(('School', 'City', 'Region', 'Applied', 'Admitted', 'Enrolled')) # schools doesn't change
True
```
```
>>> sorted(schools_admit.labels) == sorted(('School', 'City', 'Region', 'Applied', 'Admitted', 'Enrolled', 'Admission Rate')) # schools_admit doesn't change
True
```
```
>>> subset = schools_yield.where("School", are.contained_in(["21ST CENTURY EXPERIMENTAL SCH", "ABINGTON HEIGHTS HIGH SCHOOL"]))
>>> all(np.isclose(subset.column("Admission Rate"), [0.832, 0.8]))
True
```
```
>>> subset = schools_yield.where("School", are.contained_in(["21ST CENTURY EXPERIMENTAL SCH", "ABINGTON HEIGHTS HIGH SCHOOL"]))
>>> all(np.isclose(subset.column("Yield Rate"), [0.317308, 0]))
True
```


Tests for q5b_ii (points: 1)
```
>>> isinstance(is_admit_yield_correlated, bool)
True
```


Tests for q6b (points: 1)
```
>>> isinstance(overlooked_factors,np.ndarray)
True
```
```
>>> all([x in [1, 2, 3, 4] for x in overlooked_factors])
True
```
```
>>> len(overlooked_factors) == 1
True
```
```
>>> overlooked_factors.item(0) == 1
True
```


Tests for q7a (points: 2)
```
>>> sorted(admit_numbers.labels) == sorted(("Sex", "Admit", "Deny"))
True
```
```
>>> admit_numbers.num_rows == 2
True
```
```
>>> admit_numbers.where("Sex", "F").column("Admit").item(0) == 1494
True
```
```
>>> admit_numbers.where("Sex", "M").column("Deny").item(0) == 4704
True
```


Tests for q7b (points: 2)
```
>>> sorted(admit_rates.labels) == sorted(("Sex", "Admit", "Deny", "Applied", "Admission Rate"))
True
```
```
>>> admit_rates.num_rows == 2
True
```
```
>>> admit_rates.where("Sex", "F").column("Admit").item(0) == 1494
True
```
```
>>> admit_rates.where("Sex", "M").column("Applied").item(0) == 8442
True
```
```
>>> np.isclose(admit_rates.where("Sex", "F").column("Admission Rate").item(0), 0.345753)
True
```
```
>>> np.isclose(admit_rates.where("Sex", "M").column("Admission Rate").item(0), 0.442786)
True
```


Tests for q7d (points: 1)
```
>>> isinstance(bar_observation, int)
True
```


Tests for q8a_i (points: 1)
```
>>> 0 < admission_rate_overall < 1
True
```
```
>>> np.isclose(np.round(admission_rate_overall, 2), 0.41)
True
```


Tests for q8a_ii (points: 1)
```
>>> isinstance(female_admit_compare, str)
True
```
```
>>> female_admit_compare in ["lower", "higher", "comparable"]
True
```


Tests for q8b (points: 4)
```
>>> table1.num_rows == 2
True
```
```
>>> table1.num_columns == 5
True
```
```
>>> sorted(table1.labels) == sorted(("Sex", "Admit", "Deny", "Admit Expected", "Deny Expected"))
True
```
```
>>> table1.where("Sex", "F").column("Admit").item(0) == 1494
True
```


Tests for q9a (points: 2)
```
>>> sorted(applicants_by_dept.labels) == sorted(("Department", "F Applied", "M Applied"))
True
```
```
>>> applicants_by_dept.num_rows == 7
True
```
```
>>> applicants_by_dept.where("Department", "A").column("F Applied").item(0) == 108
True
```
```
>>> applicants_by_dept.where("Department", "B").column("M Applied").item(0) == 560
True
```


Tests for q9b (points: 2)
```
>>> sorted(admits_by_dept.labels) == sorted(("Department", "F Admitted", "M Admitted"))
True
```
```
>>> admits_by_dept.num_rows == 7
True
```
```
>>> admits_by_dept.where("Department", "A").column("F Admitted").item(0) == 89
True
```
```
>>> admits_by_dept.where("Department", "B").column("M Admitted").item(0) == 353
True
```


Tests for q9c (points: 2)
```
>>> numbers_by_dept.num_rows == 7
True
```
```
>>> numbers_by_dept.num_columns == 5
True
```
```
>>> sorted(numbers_by_dept.labels) == sorted(("Department", "F Admitted", "M Admitted"))
False
```
```
>>> np.all(numbers_by_dept.column("F Admitted") == admits_by_dept.column("F Admitted"))
True
```
```
>>> np.all(numbers_by_dept.column("M Applied") == applicants_by_dept.column("M Applied"))
True
```


Tests for q9d (points: 2)
```
>>> admit_rates_by_dept.num_rows == 7
True
```
```
>>> admit_rates_by_dept.num_columns == 3
True
```
```
>>> sorted(admit_rates_by_dept.labels) == sorted(("Department", "F Admission Rate", "M Admission Rate"))
True
```
```
>>> f_admission_rate_temp = numbers_by_dept.column("F Admitted")/numbers_by_dept.column("F Applied")
>>> np.all(np.isclose(admit_rates_by_dept.column("F Admission Rate"), f_admission_rate_temp))
True
```
```
>>> m_admission_rate_temp = numbers_by_dept.column("M Admitted")/numbers_by_dept.column("M Applied")
>>> np.all(np.isclose(admit_rates_by_dept.column("M Admission Rate"), m_admission_rate_temp))
True
```
