```python
# Initialize Otter
import otter
grader = otter.Notebook("hw01.ipynb")
```

# [Functions] Homework 1 – Functions and Conditionals

## Data 6, Fall 2025

In this homework assignment, you will apply what you've learned about `if`-statements, `while` loops, string manipulation, and `for` loops.

This homework is due on **Thursday, October 30 at 8:00PM**. 

You must submit the assignment to [Gradescope](https://www.gradescope.com/). See the [syllabus](https://data6.org/fa25/syllabus) for our homework late submission policy.

**Reminder:** This homework assignment contains hidden tests, which means that even if the test passed, **it does not mean your answer is 100% correct**. The autograder in the homework is intended as a sanity check.

While we encourage you to collaborate with your peers, directly copying solutions is not allowed. Please refer to our [syllabus](https://data6.org/fa25/syllabus) on Academic Honesty.
If you find yourself stuck on a problem, we recommend that you make a post on [Ed](https://edstem.org) or attend [office hours](https://data6.org/fa25/schedule/).

```python
# just run this cell
from datascience import *
import numpy as np
from IPython.display import YouTubeVideo, HTML, display
from ipywidgets import interact, widgets
%matplotlib inline
import matplotlib.pyplot as plots
plots.style.use('fivethirtyeight')
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />


# Part 1: Control Flow

<br></br>
<hr style="border: 1px solid #fdb515;" />

# Question 1 – Short-Circuit

Your friend Edwin has defined the function `some_equal`, which takes in three arguments, `a`, `b`, and `c`. Edwin wants `some_equal` to return `True` if *at least two* of `a`, `b`, and `c` are equal to one another, and `False` otherwise; his implementation is defined in the following cell:

```python
def some_equal(a, b, c):
    return a == b or b == c or a == c or 1/0
```

Sometimes, `some_equal` works:

```python
some_equal('hey', 'hi', 'hey')
```

```python
some_equal(3, 3, 3)
```

But other times...

```python
some_equal(1, 2, 4)
```

<!-- BEGIN QUESTION -->

Edwin shows you his code and asks:
> "Why don't I get a `division by zero` error in the first two examples?"

Explain to Edwin what is happening, **in your own words**. Write your answer in the cell below.

_Hint_: See lecture slides. Don't hesitate to ask on Ed or in office hours if you're not sure!


_Type your answer here, replacing this text._

<!-- END QUESTION -->

<br></br>
<hr style="border: 1px solid #fdb515;" />

# Question 2 – Death and Taxes

The United States, like many countries, uses a progressive tax bracket system. This means that as your earnings increase, the percentage of your earnings you owe in tax also increases. In addition, the US tax system uses **marginal tax brackets** – what this means is that US taxpayers pay different tax percentages on different "chunks" of their earnings. 
To read more about the US federal income tax system and marginal tax brackets, see the [Internal Revenue Service page](https://www.irs.gov/filing/federal-income-tax-rates-and-brackets). We discuss the main details below.

Let's suppose the tax brackets for this tax year are defined by the table below. This is pretty close to the actual bracket, but for simplicity we'll use 4 brackets instead of 7.

| Tax Rate | on Taxable Income from ... | up to ... | 
| --- | --- | --- | 
| 10% | $\$$0 | $\$$11,600 |
| 12% | $\$$11,601 | $\$$47,150 |
| 22% | $\$$47,151 | $\$$100,525 |
| 24% | $\$$100,526 | and up |

**You will need to use these numbers throughout this question.**

A few notes:
- We will assume all incomes are integers.
- **Taxable Income** refers to the part of ones income that is taxable; in the US there is what's known as a "standard deduction" which can be thought of as a discount on your taxes. In this assignment, we won't worry about deductions.

If someone has a taxable income of $\$$60,000, we say they are in the 22 $\%$ tax bracket. This does not mean that the individual owes 22 $\%$ of \$60,000 in taxes!

Instead, marginal tax brackets means that the individual owes the sum of the taxes for each portion of their taxable income in each tax bracket:
- 10$\%$ of $\$$11,600 = $\$$ 1160
- 12$\%$ of ($\$$47,150 - $\$$ 11,600) = 12$\%$ of $\$$35,550 = $\$$4,266
- 22$\%$ of ( $\$$60,000 - $\$$47,150) = 22$\%$ of $\$$12,850 = $\$$2,827

for a total of $\$$8,253. This makes their **effective tax rate** approximately 13.75$\%$ of $\$$60,000.

---
## Question 2a

Complete the implementation of the function `tax_rate`, which takes in a taxable income as a number (`income`) and returns the rate it is taxed at, as a fraction of 1. For instance, `tax_bracket(60000)` should evaluate to `0.22` and `tax_bracket(402150)` should evaluate to `0.24`.

_Hint_: Use what you know about conditional statements (`if-elif-else`) blocks.


```python
def tax_bracket(income):
    if ...:
        ...
    elif ...:
        ...
    elif:
        ...
    else:
    ...
```

```python
grader.check("q2a")
```

---
## Question 2b

Complete the implementation of `compute_tax`, which takes in a taxable income (`income`) and returns the amount of tax owed by an individual with that taxable income, using the marginal tax brackets described above. For instance, `tax_owed(60000)` should evaluate to `8253.0`.

**Hints**:
* There are many approaches to this problem. A straight-forward strategy would be to use one large `if-elif-...elif-else` expression, which may lead to a very long code. This approach is totally fine.
* A more organized but trickier strategy:
    * To avoid long numeric expressions with repeated computation, we suggest keeping track of several values: `total_tax` and `remaining_income`.
    * Then, you can compute the amount of income  in each tax bracket (`income_taxed`) with which to compute the tax incurred for that tax bracket.
    * The amount of income taxed at each tax bracket is the lesser of `remaining_income` and the range of the tax bracket. `min` may be useful here.
    * If there is insufficient `remaining_income` to move onto the next tax bracket, `return` early.

Regardless of the approach you choose, here's the tax bracket table again for your convenience:

| Tax Rate | on Taxable Income from ... | up to ... | 
| --- | --- | --- | 
| 10% | $\$$0 | $\$$11,600 |
| 12% | $\$$11,601 | $\$$47,150 |
| 22% | $\$$47,151 | $\$$100,525 |
| 24% | $\$$100,526 | and up |

```python
def compute_tax(income):
    ...

# the below line is for testing
compute_tax(60000)
```

```python
grader.check("q2b")
```

---
## Question 2c

Finally, complete the implementation of `effective_tax_rate`, which takes in a taxable income (`income`) and returns the effective tax rate for an individual with that taxable income, as a decimal. For instance, `effective_tax_rate(60000)` should evaluate to approximately `0.13755`.

*Note*: If `income` is 0, your `effective_tax_rate` function should also return 0. Make sure you handle this case in your function!

_Hint_: You should call the `compute_tax` function you wrote in the previous part. Our entire solution is only three lines, but you may use more than that if necessary.

```python
def effective_tax_rate(income):
    ...

# the below line is for testing
effective_tax_rate(60000)
```

```python
grader.check("q2c")
```

---
## Fun Demo

You just did all of the math necessary to build an income tax calculator, like the ones you see on many financial websites. Run the following cell once you've completed the rest of this question!

```python
# Don't worry about the code, just play with the slider that appears after running.
def plot_income_calculator(income):
    slider = f'''
    <div style="font-family: Arial, sans-serif; text-align: center;">
        <div style="margin-bottom: 20px;">
            <h1 style="font-size: 36px; color: #80B2D9;">${'{:,.2f}'.format(income)}</h1>
            <h3 style="color: #6c757d;">Taxable Income</h3>
        </div>
        
        <div style="margin-bottom: 20px;">
            <h1 style="font-size: 36px; color: #AAE895;">${'{:,.2f}'.format(compute_tax(income))}</h1>
            <h3 style="color: #6c757d;">Tax Owed</h3>
        </div>
        
        <div>
            <h1 style="font-size: 36px; color: #F99D81;">{'{:,.1f}'.format((100 * effective_tax_rate(income))) + '%'}</h1>
            <h3 style="color: #6c757d;">Effective Tax Rate</h3>
        </div>
    </div>
    '''
    
    display(HTML(slider))
    
interact(plot_income_calculator, income=widgets.IntSlider(min=0, max=200000, step=500, value=100000));
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

# Part 2: #nofilter

In this part, you will write code that applies greyscale, sepia, and x-ray filters to images! Run the below two cells to get started.

```python
# just run this cell
!pip install scikit-image
```

```python
# Don't worry about what this code is doing, just run it.
import skimage.io as skio
import matplotlib.pyplot as plt

def show_image(image_array):
    plt.figure(figsize = (10, 5))
    plt.imshow(image_array)
    plt.axis('off');
    
campanile = skio.imread('campanile.jpg')
show_image(campanile)
print('Our beautiful campus.')
```

Images are composed of **pixels**. One of the most common ways to represent the color of a pixel is using the RGB format. In RGB, each color is described by three numbers – a red (R) value, a green (G) value, and a blue (B) value – each describing the "amount" of that base color that is present. The larger R is, the more red there is.

R, G, and B are each numbers between 0 and 255 inclusive, and we often express them together in a list or array (in the order R, then G, then B). So for example, the pixel `array([255, 0, 0])` is the color red, and the pixel `array([92, 45, 99])`  is a dark-ish purple. This tool) helps visualize what different RGB values look like.


<img src="rgbmodel.jpg" style="width: 70%"/>

One way to apply filters to images is to apply mathematical functions to R, G, and B to calculate a new color for each pixel. That's exactly what you'll do in this question.

Note: We have provided a lot of skeleton code to help you display images. You only have to fill in ~7 lines total.

<br></br>
<hr style="border: 1px solid #fdb515;" />

# Question 3

---

## Question 3a: Grayscale Filter

To convert a pixel to grayscale, average its R, G, and B values:

$$
\begin{align*}
\text{new R} &= \frac{\text{old R} + \text{old G} + \text{old B}}{3} \\
\text{new G} &= \text{new B} = \text{new R}
\end{align*}
$$

That is, to convert a pixel to greyscale, you average its R, G, and B values.

Below, complete the implementation of `grayscale_filter`, which takes in a pixel (i.e., a 3-element-long NumPy array) and returns a grayscale pixel (again, a 3-element-long NumPy array).array. Example behavior:

```python
>>> grayscale_filter(make_array(50, 100, 150))
array([100., 100., 100.])

>>> grayscale_filter(make_array(132, 250, 100))
array([160.66666667, 160.66666667, 160.66666667])
```

_Hint_: Use `np.mean` to average values.

```python
def grayscale_filter(pixel):
    """Input: pixel is a 3-element NumPy array.
    Output: the pixel, in grayscale
    """
    # Don't add any more lines!
    average = ...
    grayscale_pixel = ...
    return grayscale_pixel


# for testing
grayscale_filter(make_array(50, 100, 150))
```

```python
grader.check("q3a")
```

Now that you've implemented your first filter function, it's time to see it in action. Run the cell below to convert our Campanile image to greyscale and display it.

```python
# just run this cell; it may take a few seconds
def apply_filter_to_image(image, filter_fn):
    new_image = np.zeros(image.shape)
    for x in range(image.shape[0]):
        for y in range(image.shape[1]):
            new_image[x][y] = filter_fn(image[x][y])
    
    return new_image.astype(int)

campanile_grayscale = apply_filter_to_image(campanile, grayscale_filter)

show_image(campanile_grayscale)
```

Cool!

---

## Question 3b: X-Ray Filter

Another fun filter is the x-ray filter, which inverts grayscale:

$$
\begin{align*}
\text{new R} &= 255 - \frac{\text{old R} + \text{old G} + \text{old B}}{3} \\
\text{new G} &= \text{new B} = \text{new R}
\end{align*}
$$


Below, complete the implementation of `xray_filter`, which takes in a `pixel` and returns a new pixel whose values are converted to x-ray. Example behavior is shown below.

```python
>>> xray_filter(make_array(50, 100, 150))
array([155., 155., 155.])

>>> xray_filter(make_array(132, 250, 100))
array([94.33333333, 94.33333333, 94.33333333])
```

**Note:** `xray_filter` can be written in terms of `grayscale`!

```python
def xray_filter(pixel):
    """Input: pixel is a 3-element NumPy array.
    Output: the pixel, filtered via an x-Ray
    """
    # Don't add any more lines!
    grayscale = ...
    xray_pixel = ...
    return xray_pixel

# for testing
xray_filter(make_array(50, 100, 150))
```

```python
grader.check("q3b")
```

Run the below cell to see the x-ray filter:

```python
# just run this cell; it may take a few seconds
campanile_xray = apply_filter_to_image(campanile, xray_filter)
show_image(campanile_xray)
```

<hr style="border: 1px solid #fdb515;" />

## [Tutorial] `np.dot` and `np.clip`

To implement our last filter (sepia), we need two additional NumPy functions as below. You are **not** expected to memorize these two functions for any exam; instead, what we might expect is that you are comfortable enough reading documentation about new functions to try them out.

The `np.dot` function takes in two arrays, multiplies their elements pairwise, and sums together the result. If `x` and `y` are two arrays, `np.dot(x, y)` calculates `x[0]*y[0] + x[1]*y[1] + x[2]*y[2] + ...`. Example:

```python
# just run this cell
a = make_array(5, 1, 2)
b = make_array(1, -1, 4)

# Equal to 5*1 + 1*(-1) + 2 * 4
np.dot(a, b)
```

The `np.clip` function takes in three arguments: an array, a lower bound, and an upper bound. It returns a new array where every element less than the lower bound is replaced with the lower bound, and every element greater than the upper bound is replaced with the upper bound. Example, where we clip `numbers` to be bound between 0 and 20:

```python
# just run this cell
numbers = make_array(15, -1, 25, 18, 26, 4, 5, -4)

np.clip(numbers, 0, 20)
```

---

## Question 3c: Sepia Filter

The last filter we will apply is the [sepia](https://raw.pics.io/sepia-image) filter, which will give our images a "historic" feel. The sepia filter follows this more complicated rule:

$$
\begin{align*}
\text{new R} &= (0.393 \cdot \text{old R}) + (0.769 \cdot \text{old G}) + (0.189 \cdot \text{old B}) \\
\text{new G} &= (0.349 \cdot \text{old R}) + (0.686 \cdot \text{old G}) + (0.168 \cdot \text{old B}) \\
\text{new B} &= (0.272 \cdot \text{old R}) + (0.534 \cdot \text{old G}) + (0.131 \cdot \text{old B})
\end{align*}
$$

If any of $\text{new R}$, $\text{new G}$, or $\text{new B}$ are greater than 255, we set them equal to 255.


Below, complete the implementation of `sepia_filter`, which takes in a `pixel` and returns a sepia pixel. Example behavior is shown below.

```python
>>> sepia_filter(make_array(50, 100, 150))
array([124.9 , 111.25,  86.65])

>>> sepia_filter(make_array(132, 250, 100))
array([255, 234.368, 182.504])
```

_Hints_: Use `np.dot` and `np.clip`. We have defined `red_coef`, `green_coef`, and `blue_coef` for your convenience. You may need more lines than provided.

```python
def sepia_filter(pixel):
    """Input: pixel is a 3-element NumPy array.
    Output: the pixel, sepia-filtered
    """
    # for your convenience
    red_coef = make_array(0.393, 0.769, 0.189)
    green_coef = make_array(0.349, 0.686, 0.168)
    blue_coef = make_array(0.272, 0.534, 0.131)
    
    # you may need more lines than provided.
    sepia_pixel = ...
    return sepia_pixel

# for testing
sepia_filter(make_array(50, 100, 150))
```

```python
grader.check("q3c")
```

Let's enjoy the fruits of our labor once more:

```python
# just run this cell; it may take a few seconds
campanile_sepia = apply_filter_to_image(campanile, sepia_filter)
show_image(campanile_sepia)
```

<!-- BEGIN QUESTION -->

---
## Question 3d: Custom images

Now that you've done all of the work necessary to implement three cool filters, it's only right that you get the chance to try them out on your own images. This question is graded on completeness.

**First**, run this cell. It will show you an "Upload" button which you can click to upload an image from your computer. **You must upload a `.jpg` image!**

The below cell will also show you a dropdown menu which will allow you to select from one of the three filters we created in this question. (Smaller images work faster.) **Select one of the filters from the dropdown.**

```python
# run this cell, then upload your image
def extract_array_from_upload(uploader):
    content = list(uploader.value)[0]['content']
    img = Image.open(io.BytesIO(content))
    arr = np.asarray(img)
    return arr

uploader = widgets.FileUpload(
    accept='.jpg',
    multiple=False
)

display(uploader)

filter_select = widgets.Dropdown(
    options=['grayscale', 'x-ray', 'sepia'],
    value='grayscale'
)
display(filter_select)
```

Next, run the following cell to display your uploaded image. 

```python
# after uploading your image, run this cell

import io
from PIL import Image

original_image = extract_array_from_upload(uploader)
show_image(original_image)
plt.show()
```

Finally, run the following cell to perform all of the necessary computation and output your filtered image.

```python
# just run this cell. It may take a few seconds.
filter_map = {'grayscale': grayscale_filter, 'x-ray': xray_filter, 'sepia': sepia_filter}
image_converted = (apply_filter_to_image(original_image, filter_map[filter_select.value]))
print('Filtered image')
show_image(image_converted)
```

Who needs Instagram? 😎

<!-- END QUESTION -->

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />


# Part 3: Custom Table Filters

With the power of functions and conditional statements—and your existing powers of table methods for transforming tables—you can do many, many types of data analysis. This part is lengthy, but we hope it is exciting...!

Make sure to stop by office hours or ask questions on Ed if you need assistance with any part of this question! You have come very far in your Python journey, and we are with you every step of the way.

<br></br>
<hr style="border: 1px solid #fdb515;" />

# Question 4: 2024 Cal Football Season

We would like to understand how the Cal Football team did in the 2024 season. A football game is divided into four periods, called **quarters**. The number of points Cal scored in each quarter and the number of points their opponent scored in each quarter are stored in a table called `cal_football_2024.csv`. The games are in date order.

```python
# just run this cell
football = Table().read_table("cal_football_2024.csv")
football
```

---

## Question 4a

Let's simplify team names. Assign `football_opp` to the `football` table with an added column, `Team`. This new column should contain only the strings `"Opp"` (for Opponent) or `"Cal"` (for, well, Cal), depending on the corresponding `Team Name` value. Here's an example of the first two rows of the new `football_opp` table:

| Game | Team Name | Q1 | Q2 | Q3 | Q4 | Team |
| --- | --- | --- | --- | --- | --- | --- |
| 1 | UC Davis | 6 | 7 | 0 | 0 | Opp |
| 2 | Cal | 7 | 7 | 10 | 7 | Cal |
| ... |

_Hints/Strategy_: You may want to try the following strategy, which is outlined for you in the code below:
  
1. Write a function `opp_or_cal` that takes in a team name and returns the string `"Opp"` or `"Cal"`, based on the team name. (Use conditional statements!)
1. `apply` this function to the `football` table using the `Team Name` column and get an array of strings; assign this to `team_arr`.
1. Add `team_arr` to the `football` table in a column called `Team`, and assign this to `football_opp`.

```python
def opp_or_cal(team_name):
    if ...:
        ...
    else:
        ...

team_arr = ...
football_opp = ...
football_opp
```

```python
grader.check("q4a")
```

<hr style="border: 1px solid #fdb515;" />

## [Tutorial] Row objects

We can get specific row objects from a table. You can use `tbl.row(n)` to get the `n`th row of a table. `row.item("column_name")` will allow you to select the element that corresponds to `column_name` in a particular row. Here's an example:

```python
# just run this cell
# Take off that red shirt!
football.row(20)            # <-- this will return a row object
```

```python
# just run this cell
# We got the Axe!
football.row(21)            # <-- this will return a different row object
```

```python
# just run this cell
football.row(21).item("Q4") # <-- this will return a item (e.g. an int) from a row object
```


In the next two parts of this question, we want to compute each team's score.

---

## Question 4b
Write a function `get_total`.
It should take one argument: a row object from the `football_opp` table. It should return the team's total score for that game as an integer.

**Hints/notes**:
* "Row object" means a row from the table extracted (behind the scenes) using `tbl.row(index)` that contains all the data for that specific row. It is **not** the index of a row. Do not try and call `football_opp.row(row)` inside of the function.
* If you're still confused by row objects, try `print`ing out `football_opp.row(21)` in a new cell to visually see what it looks like, or check out the tutorial above! This piece of code is pulling out the row object located at index 1 of the `football_opp` table and returning it. When you display it in a cell, you'll see that it is not located within a table, but is instead a standalone row object!
* If item from a row object are integers (like each quarter's score), you can get the items and them sum them up using imteger operators like `+`.

```python
def get_total(row):
    ...

# do not change this line
get_total(football_opp.row(21))
```

```python
grader.check("q4b")
```

---

## Question 4c

Assign `football_scores` to the `football_opp` table with an added column, `Total`. This new column should contain the team's total score for that game. Here's an example of the first two rows of the new `football_scores` table:

| Game | Team Name | Q1 | Q2 | Q3 | Q4 | Team | Total |
| --- | --- | --- | --- | --- | --- | --- | --- | 
| 1 | UC Davis | 6 | 7 | 0 | 0 | Opp | 13 |
| 2 | Cal | 7 | 7 | 10 | 7 | Cal | 31 |
| ... |

_Hint_: Use your newly defined `get_total` function, and check out the `apply` documentation in the Data 6 Python Reference.

```python
football_scores = ...
football_scores
```

```python
grader.check("q4c")
```

In the final two parts of this question, we want to count how many wins Cal had.

---

## Question 4d

Use `football_scores` to create `football_games`, a new table of game scores between Cal and each opponent. The first few rows of your table will look like this:

| Game | Cal | Opp |
| --- | --- | --- |
| 1 | 31 | 13 |
| 2 | 21 | 14 |
| ... |

_Hints_:

* You will likely use the `pivot` table method. Refer to the [Data 6 Python Reference](https://data6.org/notes/reference) if you're unsure.
* The `collect` argument to `pivot`: There is exactly one score for each entry in the pivot table, e.g., Cal has exactly one score for Game 1: 31. How do you get the value of a single-length array?
* Cal played 13 games in the 2024 football season, so your `football_games` table should have 13 rows.

```python
football_games = ...
football_games.show(13)
```

```python
grader.check("q5a_ii")
```

---

## Question 4e

Finally, we want to see how Cal did this season. Using the `football_games` table:

1. Make a new table, `football_results, that adds a column named `Win` to the `football_games` table. This column will be a series of `True` and `False` values corresponding to whether or not Cal won that game.
1. Then, assign the number of wins and losses Cal had to `cal_wins` and `cal_losses`, respectively.

```python
football_results = ...

cal_wins = ...
cal_losses = ...

# Don't delete or edit the following line:
display(football_results)
print(f"In the 2024 Season, Cal Football won {cal_wins} games and lost {cal_losses} games. Go Bears! 🐻")
```

```python
grader.check("q4e")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />


# Part 4: Strings, Slices, and (briefly) Lists

Let's practice some string methods that we have briefly introduced in the lecture notes. 

We have already seen strings before in this course, but we want to introduce some new ideas with strings that will be very useful when dealing with data from real-world data sets.

<hr style="border: 1px solid #fdb515;" />

# Question 5

### [Tutorial] Indexing and Slicing

Strings can be viewed as special sequences of characters. Just like arrays, they can be **indexed** and **sliced**. We define each of these features below.

```python
# just run this cell
our_string = "Data Science at Berkeley rocks!"
our_string
```

**Indexing** into a string is getting the character at a particular index of the string. Strings are zero-indexed.

Unlike arrays, strings do not have an `item` method. However, we can use **square bracket notation** to index into strings:

```python
our_string[0] + our_string[5]
```

We can also use negative indexing:

```python
our_string[-1]
```

**Slicing** into a string, or getting a **slice** of a string, is getting a sequence of characters between two indexes. Slices also use square brackets, but now we specify the starting and ending indices of the range:

```python
s[start:end]
```

Similar to array ranges (e.g., `np.arange`), the starting index is included in the slice, and the ending index is excluded from the slice.

```python
our_string[16:24]
```

**Slicing, full square bracket notation**:
While slices by default will take every character in the `[start, end)` range of indices, square bracket notation can also specify the `step`, or number of characters to skip. This behavior is just like array ranges.

```python
s[start:end:step]
```

```python
our_string[1:-2:2]
```

If `start` is not specified, then the default value of `start` is assumed to be 0. this behavior is just like array ranges:

```python
our_string[:4]
```

If `end` is not specified, then the default value of `end` is assumed to be the length of the string, such that the slice runs to the end of the string. This differs from array ranges, which must have an end:

```python
our_string[-6:]
```



If `step` is not specified, the default value of `step` is assumed to be 1. This behavior is just like array ranges.

```python
our_string[1:-2]
```

All parts of this question will use the following demo string. This string of characters is very common in typesetting and typing examples, because it is a sentence that contains all the letters of the alphabet. Read more on [Wikipedia](https://en.wikipedia.org/wiki/The_quick_brown_fox_jumps_over_the_lazy_dog).

```python
# just run this cell
demo_str = "The quick brown fox jumps over the lazy dog"
demo_str
```

---

### Question 5a

Assign `q5a_start` and `q5a_end` to integers such that when the cell is run, it outputs ``quick``.

```python
q5a_start = ...
q5a_end = ...

# do not edit below this lien
demo_str[q5a_start:q5a_end]
```

```python
grader.check("q5a")
```

<!-- BEGIN QUESTION -->

---

### Question 5b

Assign `q5b_slice` to a slice of `demo_str` such that when the cell is run, it prints

```
qkofjsehad
```

_Hint_: Remember `s[start:end:step]`


This question is manually graded.

```python
q5b_slice = ...

# do not edit below this line
print(q5b_slice)
```

<!-- END QUESTION -->

<hr style="border: 1px solid #fdb515;" />

# Question 6

## [Tutorial] Lists and slices

Lists are a built-in Python data type that we won't get into much into this course, but necessarily you will interact with when working with Python.

For example, the `split` string method divides a string into substrings (smaller strings). `split` returns a list, denoted by square brackets:

```python
our_string.split(' ')
```

```python
type(our_string.split(' '))
```

**Lists are very similar to arrays**; the precise details are covered in our lecture notes. In this tutorial, we will just describe the **square-bracket indexing** common to lists and to strings.

```python
our_string.split(' ')[-2]
```

```python
our_string.split(' ')[1:3]
```

All arrays can be converted into lists. Run the two cells below, and note the difference in notation:

```python
arr = make_array("Go", "Bears", "!")
arr
```

```python
l = list(arr)
l
```

We can use this array-to-list conversion to pass in arrays to functions that may have list or list-like data types as arguments. For example, the `join` string method concatenates the elements in a list, separated by the provided string itself:

```python
'...'.join(arr)
```

For this question, we will continue using our typesetting demo string for just a bit more:

```python
# just run this cell
demo_str
```

<!-- BEGIN QUESTION -->

---

## Question 6a

Assign `q6a_str` to an expression involving `demo_str` such that when the below cell is run, it prints

```THE,QUICK,BROWN,FOX,JUMPS,OVER,THE,LAZY,DOG```

_Hint_: Use string methods—you might want to look at `join` and `split` and how they are used in the [apply examples](https://data6.org/notes/13-functions/apply.html#applying-with-strings). This question is manually graded.

```python
q6a_str = ...

# do not edit below this line
print(q6a_str)
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

---

## Question 6b

Assign `q6b_str` to an expression involving `demo_str` such that when the below cell is run, it prints

```
THE
QUICK
BROWN
FOX
```

_Hint_: Recall the special newline character, `'\n'`. This question is manually graded.

```python
q6b_str = ...

# do not edit below this line
print(q6b_str)
```

<!-- END QUESTION -->

# Done!

Congrats!

The point breakdown for this assignment is given in the table below:

| **Category** | Points |
| --- | --- |
| Autograder | 32.5 |
| Written  | 5 |
| **Total** | 37.5 |

## Pets of Data 6

Yoshi said "Purrr" for finishing this homework!

<img src="yoshi.jpeg" width="50%" alt="Cute Yoshi"/>

## Submission

Below, you will see two cells. Running the first cell will automatically generate a PDF of all questions that need to be manually graded, and running the second cell will automatically generate a zip with your autograded answers. You are responsible for submitting both the coding portion (the zip) and the written portion (the PDF) to their respective Gradescope portals. **Please save before exporting!**

> **Important: You must correctly assign the pages of your PDF after you submit to the correct gradescope assignment. If your pages are not correctly assigned and/or not in the correct PDF format by the deadline, we reserve the right to award no points for your written work.**

If there are issues with automatically generating the PDF in the first cell, you can try downloading the notebook as a PDF by clicking on `File -> Save and Export Notebook As... -> PDF`. If that doesn't work either, you can manually take screenshots of your answers to the manually graded questions and submit those. Either way, **you are responsible for ensuring your submission follows our requirements, we will NOT be granting regrade requests for submissions that don't follow instructions.**

```python
from otter.export import export_notebook
from os import path
from IPython.display import display, HTML
name = 'hw04'
export_notebook(f"{name}.ipynb", filtering=True, pagebreaks=True)
if(path.exists(f'{name}.pdf')):
    display(HTML(f"Download your PDF <a href='{name}.pdf' download>here</a>."))
else:
    print("\n Pdf generation failed, please try the other methods described above")
```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q2a (points: 3)
```
>>> callable(tax_bracket)
True
```
```
>>> tax_bracket(60000) == 0.22
True
```
```
>>> tax_bracket(402150) == 0.24
True
```
```
>>> tax_bracket(40125) == 0.12
True
```
```
>>> tax_bracket(1) == 0.1
True
```


Tests for q2b (points: 4)
```
>>> callable(compute_tax)
True
```
```
>>> np.isclose(compute_tax(150), 15.0)
True
```
```
>>> np.isclose(compute_tax(60000), 8253.0)
True
```


Tests for q2c (points: 2)
```
>>> np.isclose(effective_tax_rate(60000), 0.13755)
True
```
```
>>> np.isclose(effective_tax_rate(12000), 0.10066666666666667)
True
```


Tests for q3a (points: 3)
```
>>> # This test checks the first example in question description.
>>> np.all(np.isclose(grayscale_filter(make_array(50, 100, 150)), make_array(100, 100, 100)))
True
```
```
>>> # This test checks the second example in question description.
>>> np.all(np.isclose(grayscale_filter(make_array(132, 250, 100)),
...                   make_array(160.66666667, 160.66666667, 160.66666667)))
True
```


Tests for q3b (points: 3)
```
>>> # This test checks the first example in question description.
>>> np.all(np.isclose(xray_filter(make_array(50, 100, 150)), make_array(155, 155, 155)))
True
```
```
>>> # This test checks the second example in question description.
>>> np.all(np.isclose(xray_filter(make_array(132, 250, 100)),
...                   make_array(94.33333333, 94.33333333, 94.33333333)))
True
```


Tests for q3c (points: 3)
```
>>> # This test checks the first example in question description.
>>> np.all(np.isclose(sepia_filter(make_array(50, 100, 150)), make_array(124.9 , 111.25,  86.65)))
True
```
```
>>> # This test checks the second example in question description.
>>> np.all(np.isclose(sepia_filter(make_array(132, 250, 100)),
...                   make_array(255, 234.368, 182.504)))
True
```


Tests for q4a (points: 3)
```
>>> football_opp.num_rows == football.num_rows
True
```
```
>>> football_opp.num_columns == football.num_columns + 1
True
```
```
>>> "Team" in football_opp.labels
True
```
```
>>> football_opp.where("Team Name", "UC Davis").row(0).item("Team") == "Opp"
True
```
```
>>> np.all(football_opp.where("Team Name", "Cal").column("Team") == "Cal")
True
```


Tests for q4b (points: 3)
```
>>> callable(get_total)
True
```
```
>>> get_total(football_opp.row(21)) == 24
True
```
```
>>> get_total(football_opp.row(20)) == 21
True
```


Tests for q4c (points: 2)
```
>>> football_scores.num_rows == football.num_rows
True
```
```
>>> football_scores.num_columns == football.num_columns + 2
True
```
```
>>> "Team" in football_scores.labels
True
```
```
>>> "Total" in football_scores.labels
True
```
```
>>> football_scores.where("Team Name", "UC Davis").row(0).item("Total") ==  13
True
```


Tests for q4e (points: 3)
```
>>> football_results.num_rows == football_games.num_rows
True
```
```
>>> football_results.num_columns == football_games.num_columns + 1
True
```
```
>>> cal_wins >= 0
True
```
```
>>> cal_losses >= 0
True
```
```
>>> np.all([x in (True, False) for x in football_results.column("Win")])
True
```


Tests for q5a (points: 0.5)
```
>>> isinstance(q5a_start, int) and isinstance(q5a_end, int)
True
```
```
>>> demo_str[q5a_start:q5a_end] == "quick"
True
```


Tests for q5a_ii (points: 3)
```
>>> football_games.num_columns == 3
True
```
```
>>> football_games.num_rows == 13
True
```
```
>>> sorted(football_games.labels) == ['Cal', 'Game', 'Opp']
True
```
