```python
# Initialize Otter
import otter
grader = otter.Notebook("hw01.ipynb")
```

# Homework 1 – Introduction to Python and Jupyter

## Data 6, Fall 2025

In this homework assignment, you will apply what you've learned about arithmetic expressions, data types, boolean expressions, and functions. Your job is to write code anywhere you see `...`, as well as answer any open-ended questions.

This homework is due on **Wednesday, September 10 at 5:00 PM**.

You must submit the assignment to Gradescope. Submission instructions can be found at the bottom of this notebook and on Ed. See the [syllabus](https://data6.org/fa25/syllabus) for our late submission policy.

```python
# Just run this cell to import necessary Python packages
try:
    from datascience import *
except:
    %pip install -q datascience
    from datascience import *
import numpy as np
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

# Part 1: Computation

<br/><br/>


<hr style="border: 1px solid #fdb515;" />

# Question 1 – Errors

In this question, you will identify why certain lines of code cause an error or have unexpected behavior.

---

## Question 1a

The following cell contains an error.

```python
8 = 10 - 2
```

Choose the best explanation of what's wrong with the code, and then assign the integer 1, 2, 3, or 4 to `reason_1a` below to indicate your answer.

1. Python already knows `8 = 10 - 2`.

2. `8` is a number, and numbers can't be Python names. In Python, "`x = 10 - 2`" means "assign `x` as the name for the value of `10 - 2`."

3. It should be `10 - 2 = 8`.

4. I don't get an error message. This is a trick question.

```python
# Your answer should be in the format reason_1a = n, where n is either 1, 2, 3, or 4

reason_1a = ...
```

```python
grader.check("q_1a")
```

---

## Question 1b

The following cell contains an error.

```python
four = 3
nine = four times four
```

Choose the best explanation of what's wrong with the code and assign the integer `1`, `2`, `3`, or `4` to `reason_1b` below to indicate your answer.

1. The `times` operation only applies to numbers, not the word `four`.

2. The name `four` cannot be assigned to the number 3.

3. four times four is sixteen, not nine.

4. The name `four` cannot be followed directly by another name.

```python
# Your answer should be in the format reason_1b = n, where n is either 1, 2, 3, or 4

reason_1b = ...
```

```python
grader.check("q_1b")
```

---

## Question 1c

Terry runs the following cell and expects that the value of `eight` will be `8`, stored as an `int`.

```python
eight = len("3" + "5")
```

What data type is the value bound to the name `eight`? In other words, does `eight` have type `int`, `float`, `str`, `bool`, or `NoneType`?

Assign your answer to the name `eight_type` below **as a string**. For example, if you think the answer is `int`, write `eight_type = 'int'` below.

```python
eight_type = ...
```

```python
grader.check("q_1c")
```

---

## Question 1d

Sam wants to create and print a string consisting of their name, age, and house address, as follows.

```python
name = 'Sam'
age = 51
house_number = 1436
house_street = 'Durant Avenue'

sam_greeting = 'My name is ' + name + ', I am ' + age + ' years old, and I live at ' + house_number + ' ' + house_street
sam_greeting
```

Output:
```
TypeError                                 Traceback (most recent call last)
Cell In[4], line 6
      3 house_number = 1436
      4 house_street = 'Durant Avenue'
----> 6 sam_greeting = 'My name is ' + name + ', I am ' + age + ' years old, and I live at ' + house_number + ' ' + house_street
      7 sam_greeting

TypeError: can only concatenate str (not "int") to str
```

However, it errors. In the cell below, assign `sam_greeting_fixed` to a fixed version of `sam_greeting` that doesn't error. Start by copying the right-hand side of the `sam_greeting` assignment statement above. Your answer must still involve `+`.

```python
sam_greeting_fixed = ...

# If your sam_greeting_fixed is correct, the following print statement should print
# 'My name is Sam, I am 51 years old, and I live at 1436 Durant Avenue'
sam_greeting_fixed
```

```python
grader.check("q_1d")
```

<br/><br/>

<hr style="border: 1px solid #fdb515;" />

# Question 2 -  Order of Operations

Consider the following Python expression:

expression = 18 // 4 + 18 % 4 * 2

Which of the following expressions is equivalent (produces the same result as `expression`)? 

1. `(18 // 4 + 18) % (4 * 2)`

2. `(18 // 4) + 18 % 4 * 2` 

3. `(18 // 4) + (18 % (4 * 2))`

4. `18 // (4 + 18 % 4 * 2)`

equivalent_expression = ...


```python
equivalent_expression = 2
```

```python
grader.check("q_2")
```

<br/><br/>

<hr style="border: 1px solid #fdb515;" />

# Question 3 -  Names

We have seen how names can bind to values. They can also bind to function names! For example, consider the following two cells, where `another_name_for_len` is bound to the `len` function:

```python
# Demo: Just run this cell
s = "this string is length 24"
len(s)
```

```python
# Demo: Just run this cell
another_name_for_len = len
another_name_for_len(s)
```

---

## Question 3a

In the below cell, assign `mystery_cast` to a function name such that when run, the last line `str(mystery_cast(5.00)) + '3'` evaluates to the string `'53'`.

**Hint**: Try casting!

```python
mystery_cast = ...
str(mystery_cast(5.00)) + '3'
```

```python
grader.check("q_3a")
```

---

## Question 3b

Consider the following Python statement:

```python
value = max( min(10, 4 + 7), 2 * 3 )
```

Below are some sub-expressions (shuffled) that Python must evaluate to assign `value` to `10`:

A. `4 + 7`

B. `2 * 3`

C. `min(10, 11)`

D. `max(10, 6)`


In what order does Python evaluate the expression? Choose the best option below and assign 1, 2, 3, or 4 to `order_3b`. (1 Point)

1. A -> B -> C -> D

2. B -> A -> C -> D

3. A -> C -> B -> D

4. B -> C -> A -> D


<!-- BEGIN QUESTION -->



```python
order_3b = ...
```

```python
grader.check("q_3b")
```

<!-- END QUESTION -->

---

## Question 3c

Assign `mystery_func` to a function such that when run, the cell outputs the `5`.

_Hint_: Review the previous parts of this question.

```python
mystery_func = ...

# do not edit any code below this line
g = min
g(mystery_func(2, g(10, 5, 15), 0), 20)
```

```python
grader.check("q_3c")
```

<!-- BEGIN QUESTION -->

<br/><br/>

<hr style="border: 1px solid #fdb515;" />

# Question 4 – Markdown

[Markdown](https://en.wikipedia.org/wiki/Markdown) isn't a programming language – it's a "markup language", which means that it is designed to format text to look nice. But like a programming language, Markdown has a set of rules (syntax) that govern how it works. Think of it like a fancy Google Docs.

Here are some of the most important rules in Markdown:

- Double asterisks (`**`) are how we make words bold. `My name is **Jordan**` renders to "My name is **Jordan**".
- Underscores (`_`) are how we make words italicized. `How are _you_?` renders to How are _you_?, and `what is **_up_**` renders to what is **_up_**.
- Hashtags (`#`, `##`, `###`) are how we create titles/headings. The more hashtags you include, the smaller the title is. The title for this question was generated using `## Question 5 – Markdown`.
- Bulleted lists are created using the (`-`) character.
- Numbered lists are created the same way they're created in a standard word processor, like Google Docs.

Markdown is important to learn because Jupyter Notebooks are often full of Markdown cells (just like this assignment). In future assignments, we will ask you to write answers to written questions using Markdown, and while the majority of your answers will be in plaintext (i.e. with none of the special characters above), learning Markdown is valuable nonetheless.

### Task

Your job is to recreate the output in the following screenshot, using the Markdown cell below. Feel free to replace "Bongo Burger" and "K-Pop" with your own favorite food and music genre, but make sure the formatting is the same.

<img src='markdown.png' width=300>

This is the first question in this homework where we will grade your work manually instead of relying on an autograder. There will be a few of these questions.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<br/><br/>

<hr style="border: 1px solid #fdb515;" />

# Question 5 — Assignment Statements

Run the following cell.

```python
x = 7
y = x // 2
x = 9
```


What is `y` after running this cell, and why? Choose the best explanation and assign 1, 2, 3, or 4 to `names_q3` below to indicate your answer. (1 Points)

1. `y` is equal to 4, because `x` was updated to 9 and `9 // 2` is 4.

2. `y` is equal to 4, because ``7 // 2`` is 4. 

3. `y` is equal to 3, because ``7 // 2`` is 3, and changing x later does not affect y.

4. `y` is equal to 3, because Python ignores the line `x = 9` since `x` was previously defined. 





```python
names_q4 = ...
```

```python
grader.check("q_5")
```

<br/><br/> 

<hr style="border: 1px solid #fdb515;" />

In the lab worksheet, we drilled down into a data visualization of traffic accidents. In particular, one question discussed the idea of the presented data as **age-standardized**, meaning that rates are calculated separately for each age group in a country and then combined according to a specific weighting, where the weighting is the same for each country. Let's explore age standardization more in this question.

The national statistics office of Canada provides a detailed explanation of how to calculate [age-standardized rates](https://www.statcan.gc.ca/en/dai/btd/asr). We will build off of this webpage, so we recommend you read it and understand the main points before continuing.

# Question 6 — Crude Rates

The article above discusses how to compare cancer mortality rates across Canada for two different years, 2000 and 2011. It defines a year's **crude rate** for cancer mortality as an incidence rate based on the population of that year. This is the same incidence rate we have seen before:

$$ \text{Crude rate} = \frac{\text{\# cases}}{\text{population}} \times 100000$$

From the article:

> [In] 2000, show that 62,672 Canadians died of cancer, while, in 2011, 72,476 died. Over this same time period, the Canadian population grew from 30,685,730 in 2000, to 34,342,780 in 2011. We see that the cancer mortality rate was 204.2 deaths per 100,000 persons in 2000, and 211.0 deaths per 100,000 persons in 2011.

However, these crude rates are **rounded**. In this question, we will reproduce these crude rates using Python and compute percent changes between the two years' crude rates.

---

## Question 6a

Using the quoted cancer mortality and population numbers in the quote above, write an expression that computes the crude rate of cancer mortality in 2000 and give the value of your expression the name `crude_2000`. Do the same for the crude rate of 2011 and give that value the name `crude_2011`.

Notes:
* You should **not** round these numbers.
* The last line displays multiple names in the same Python expression.

```python
crude_2000 = ...
crude_2011 = ...
crude_2000, crude_2011 # keep this line for display purposes
```

```python
grader.check("q_6a")
```

---

## Question 6b

What is the percent change of crude rates from 2000 to 2011? Write a **single-line** expression that computes this value to the nearest thousandth place. Your expression should use the names `crude_2000` and `crude_2011`, as well as the `round` function. Give the value of your expression the name `percent_change_crude`.

```python
percent_change_crude = ...
percent_change_crude
```

```python
grader.check("q_6b")
```

<!-- BEGIN QUESTION -->

---

## Question 6c

From the [article](https://www.statcan.gc.ca/en/dai/btd/asr):

> ...the crude rate for the total population, while it accurately represents the incidence of death due to cancer each year, it is not the correct indicator to use to compare the incidence between years.

Why? Based on the article's explanation, write a response in the cell below; your answer should not exceed 3 sentences.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<br/><br/> 

<hr style="border: 1px solid #fdb515;" />

# Question 7 — Age-Standardized Rates

An **age-standardized rate** levels the field of comparison by restructuring both populations to have the same age structure, often to a **standard population**. The article walks through this computation using a standard population from the 1991 Census to obtain:

> [the] 2000 age-standardized mortality rate of 177.9 deaths per 100,000 standard population, and
> 
> [the] 2011 age-standardized mortality rate of 163.6 deaths per 100,000 standard population.

(It is worth noting that the numbers computed in the article have significant roundoff, but the idea still stands.)

The article footnotes further mention that a "transition to a more recent age structure is being considered." Let's now consider transitioning the age standard to the more recent 2021 census.

For reference, here are the (rounded) age-specific crude rates presented in Table 1 of the article:

| Age Group | 2000 | 2011 |
| --- | --- | --- |
| 0 to 39 years | 7.9 | 5.8 |
| 40 years and over | 450.4 | 416.7 |
| Total all ages | 204.2 | 211.0 |

We have also constructed the table below, which lists the proportions of the populations in each age group for both 1991 and **2021** ([source](https://www12.statcan.gc.ca/census-recensement/2021/dp-pd/prof/details/page.cfm?Lang=E&DGUIDList=2021A000011124&GENDERList=1,2,3&STATISTICList=1,4&HEADERList=0&SearchText=Canada)).

| Age Group | 1991 | 2021 | 
| --- | --- | --- |
| 0 to 39 years | 61.6% | 47.8% |
| 40 years and over | 38.4% | 52.2% |
                                                                           
The article computes the **Age-Standardized Mortality Rate (ASMR)** as the combination of the age-specific rates, multiplied by the proportion of the standard population. 
For a 1991 standard population, these rates are (as per the article):

$$ ASMR_{2000} = (7.9 \times 61.6\%) + (450.4 \times 38.4\%) = 177.8 $$
$$ ASMR_{2011} = (5.8 \times 61.6\%) + (416.7 \times 38.4\%) = 163.6, $$

where $ASMR_{2000}$ and $ASMR_{2011}$ are the ASMR rates for 2000 and 2011, respectively, computed using the 1991 standard population proportions.

---

## Question 7a

Compute the ASMR for 2000 using the **2021** standard population proportions. Write a single-line expression computing this value and round to the nearest tenth place. Your expression can use numeric literals from the tables above. Give the value of your expression the name `asmr_2000_new`.

```python
asmr_2000_new = ...
asmr_2000_new
```

```python
grader.check("q_7a")
```

<!-- BEGIN QUESTION -->

---

## Question 7b

You will observe that the adjusted ASMR for 2000 you computed in the previous part (using a new standard population of 2021) is much closer to the crude rate for 2000, as compared to the 1991 standard population mentioned in the article. You will see the same pattern if you repeated the computation for 2011.

Briefly explain why these adjusted ASMRs for the years 2000 and 2011, computed for the standard population of 2021, are closer to the respective crude rates for these years. Your explanation should reference the ASMR equations and not any other social components beyond what is mentioned in the article. Your answer should be no longer than 2-3 sentences. 

_Hint_: Reference the full [Table 1](https://www.statcan.gc.ca/en/dai/btd/asr) in the article for the proportions of the age-specific populations in 2000 and 2011.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

## Done!

Congratulations! You've finished your first Data 6 homework assignment! To submit your work, follow the steps outlined on Ed.

| Category | Points |
| --- | --- |
| Autograder (Coding questions) | 11 |
| Written (Q3a, Q4, Q6d and Q7b) | 4 |
| **Total** | 15 |

## Pets of Data 6
Take it easy like Tigger. Congrats on finishing the first Data 6 homework!

<img src="tigger.jpg" width="50%" alt="Two cute sleepy cats on a floral cushion"/>

## Submission

Below, you will see two cells. Running the first cell will automatically generate a PDF of all questions that need to be manually graded, and running the second cell will automatically generate a zip with your autograded answers. You are responsible for submitting both the coding portion (the zip) and the written portion (the PDF) to their respective Gradescope portals. **Please save before exporting!**

> **Important: You must correctly assign the pages of your PDF after you submit to the correct gradescope assignment. If your pages are not correctly assigned and/or not in the correct PDF format by the deadline, we reserve the right to award no points for your written work.**

If there are issues with automatically generating the PDF in the first cell, you can try downloading the notebook as a PDF by colicking on `File -> Save and Export Notebook As... -> PDF`. If that doesn't work either, you can manually take screenshots of your answers to the manually graded questions and submit those. Either way, **you are responsible for ensuring your submision follows our requirements, we will NOT be granting regrade requests for submissions that don't follow instructions.**

```python
from otter.export import export_notebook
from os import path
from IPython.display import display, HTML
name = 'hw01'
export_notebook(f"{name}.ipynb", filtering=True, pagebreaks=True)
if(path.exists(f'{name}.pdf')):
    display(HTML(f"Download your PDF <a href='{name}.pdf' download>here</a>."))
else:
    print("\n Pdf generation failed, please try the other methods described above")
```

Output:
```
<IPython.core.display.HTML object>
```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Output:
```
<IPython.core.display.HTML object>
```

Tests for q_1a (points: [0, 0])
```
>>> isinstance(reason_1a, int)
True
```
```
>>> reason_1a in [1, 2, 3, 4]
True
```


Tests for q_1b (points: [0, 0])
```
>>> isinstance(reason_1b, int)
True
```
```
>>> reason_1b in [1, 2, 3, 4]
True
```


Tests for q_1c (points: [0, 0])
```
>>> isinstance(eight_type, str)
True
```
```
>>> eight_type in ['int', 'float', 'str', 'bool', 'NoneType']
True
```


Tests for q_1d (points: [1])
```
>>> isinstance(sam_greeting_fixed, str)
True
```


Tests for q_2 (points: [0.5])
```
>>> equivalent_expression in [1, 2, 3, 4]
True
```


Tests for q_3a (points: [0.5, 0.5])
```
>>> isinstance(mystery_cast, type)
True
```
```
>>> str(mystery_cast(5.00)) + '3' == '53'
True
```


Tests for q_3b (points: [0])
```
>>> 1 <= order_3b <= 4
True
```


Tests for q_3c (points: [1])
```
>>> g(mystery_func(2, g(10, 5, 15), 0), 20) == 5
True
```


Tests for q_5 (points: [0])
```
>>> 1 <= names_q4 <= 4
True
```


Tests for q_6a (points: [0, 0, 0.5, 0, 0, 0.5])
```
>>> type(crude_2000) is float
True
```
```
>>> abs(crude_2000 - round(crude_2000)) > 0 # checks unrounded
True
```
```
>>> expected_2000 = 62672/(30685730/100000)
>>> round(crude_2000, 6) == round(expected_2000, 6)
True
```
```
>>> type(crude_2011) is float
True
```
```
>>> abs(crude_2011 - round(crude_2011)) > 0 # checks unrounded
True
```
```
>>> expected_2011 = 72476/(34342780/100000)
>>> round(crude_2011, 6) == round(expected_2011, 6)
True
```


Tests for q_6b (points: [0, 0])
```
>>> type(percent_change_crude) is float
True
```
```
>>> 0 < abs(percent_change_crude) < 100
True
```


Tests for q_7a (points: [0, 0.5])
```
>>> type(asmr_2000_new) is float
True
```
```
>>> abs(asmr_2000_new - round(asmr_2000_new, 1)) == 0 # checks rounded
True
```
